```python
# Initialize Otter
import otter
grader = otter.Notebook("gla01.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Guided Learning Activity 01: Study Design

This Guided Learning Activity is designed for you to complete alongside a Data Ambassador from the course. You might find that it feels like a combination of the lectures and lab assignment. Whether you are participating live or watching the recording of the live meeting, let the Data Ambassador guide you through the following tasks. There will be moments for you to reflect and explore your own ideas as a way to solidify concepts and skills introduced by your instructor. Keep in mind that this is not a graded assignment for MATH 108 by default. If you have any concerns about participation, reach out to your instructor.

---

## Learning Objectives

1. Differentiate between experiments and observational studies, and explain their roles in establishing association and causation.
1. Navigate and edit Jupyter Notebooks to write explanatory text and execute Python code.
1. Write Python expressions, including arithmetic calculations, assignment statements, and function calls.
1. Interpret Python error messages to troubleshoot and debug code effectively.

---

## Diabetes

Diabetes is a chronic condition that occurs when the body cannot properly regulate blood sugar (glucose) levels. Glucose is the primary source of energy for the body's cells, and its levels are controlled by a hormone called insulin, produced by the pancreas. In diabetes, this system is disrupted, leading to elevated blood sugar levels (i.e. hyperglycemia), which can cause long-term health complications if left unmanaged.  

<a href="https://www.niddk.nih.gov/health-information/diabetes/overview/what-is-diabetes"><img style="display: block; margin-left: auto; margin-right: auto;" src="./BloodGlucoseImage.jpg" alt="On the left, a diagram of a blood vessel with normal blood glucose levels, containing fewer glucose molecules. On the right, a diagram of a blood vessel with high blood glucose levels, containing more glucose molecules." width="40%"></a>  

### Type 2 Diabetes Research

Diabetes can be classified into Type 1, Type 2, and other less common types. Type 2 diabetes is the most prevalent and occurs when the body either:  

- Becomes resistant to insulin, or  
- Does not produce enough insulin to maintain normal blood sugar levels.  

The [National Institute of Diabetes and Digestive and Kidney Diseases (NIDDK)](https://www.niddk.nih.gov/) and other organizations have made significant efforts to understand how to prevent and manage Type 2 diabetes.

### Study Designs

The methodology used by the researchers to obtain their results can be almost more important than the results themselves. If the wrong methodology is used, then all the work down could be completely rejected by the scientific community. You were exposed to several ideas about study design in this week's materials.

#### Task 01 💭

<!-- BEGIN QUESTION -->

What are the two fundamental study designs? What is the key distinction between the two designs?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

#### Task 02 💭

<!-- BEGIN QUESTION -->

From the 1940s through the 1970s, early observational studies identified risk factors for Type 2 diabetes, such as obesity, physical inactivity, and family history. A famous example is the Framingham Heart Study. This study primarily focused on cardiovascular disease (CVD), but it also included an analysis of Type 2 diabetes. Take a look at the Methods section (shown below) from the paper [_Diabetes and Cardiovascular Risk Factors: The Framingham Study_](./kannel-mcgee-1979-diabetes-and-cardiovascular-risk-factors-the-framingham-study.pdf) 

> The Framingham Study has been in continuous operation for 25 years, following a cohort of 5209 men and women aged 30-62 years at initial examination biennially to learn in what areas those who develop CVD differ from those who remain free of it. At each biennial examination each participant was given a thorough standardized cardiovascular examination which also elicited information concerning their physical characteristics, living habits, blood chemistry and other characteristics suspected of contributing to the occurrence of CVD. This was supplemented by information on cardiovascular illness and death from daily surveillance of hospital admissions to the only hospital in town, attending physician's and medical examiner's reports, information from the spouse and other extra-clinic sources. Details of the laboratory methods and clinical criteria used are available elsewhere. Clinical cardiovascular end points were established from biennial clinical examinations confirmed by two physician examiners, panel review of hospital protocols and death certificates, ECG findings and cardiac enzymes. A description of the method of measurement of risk factors and the precise criteria for a definition of diabetes may be found elsewhere. A person was diagnosed as having diabetes mellitus if he or she was under treatment for diabetes or had elevated casual blood glucose determinations at two successive examinations. Follow-up of the cohort has been reasonably complete, with 3% lost to follow-up for cardiovascular mortality in 20 years. For each examination there was an 85% participation level. The rest were examined at less frequent intervals, with 69% total population participation over all examinations up to and including the tenth biennial examination or until the time of their death. Few cardiovascular events were unaccounted for, since we arranged later examinations for those missed, had daily hospital surveillance and obtained other ancillary information.

After reading the description, what indicates that this was an observational study, and how does that study design affect the ability for the researchers to establish a casual relationship between the factors they explored and the risk of developing Type 2 diabetes?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

### Diabetes Prevention Program (DPP)

In the 1990s, large-scale Type 2 diabetes intervention studies began to emerge. One famous study was the The [Diabetes Prevention Program](https://www.niddk.nih.gov/about-niddk/research-areas/diabetes/diabetes-prevention-program-dpp) (DPP). This study started in 1996 and ended in 2001. Let's take a look at the major parts of that study to understand its design.

#### Task 03 💭

<!-- BEGIN QUESTION -->

The 2002 paper [_Reduction in the Incidence of Type 2 Diabetes with Lifestyle Intervention or Metformin_](./NEJMoa012512.pdf) summaries the results of the study. Review the Participants subsection of the Methods (shown below): 

> Eligible participants were randomly assigned to one of three interventions: standard lifestyle recommendations plus metformin (Glucophage) at a dose of 850 mg twice daily, standard lifestyle recommendations plus placebo twice daily, or an intensive program of lifestyle modification. The study initially included a fourth intervention, troglitazone, which was discontinued in 1998 because of the drug's potential liver toxicity. The results in the troglitazone group are not reported here. Treatment with metformin was initiated at a dose of 850 mg taken orally once a day, with placebo tablets also given once a day initially. At one month, the dose of metformin was increased to 850 mg twice daily, unless gastrointestinal symptoms warranted a longer titration period. The initiation of treatment with half a tablet was optional. Adherence to the treatment regimen was assessed quarterly on the basis of pill counts and structured interviews. The standard lifestyle recommendations for the medication groups were provided in the form of written information and in an annual 20-to-30-minute individual session that emphasized the importance of a healthy lifestyle. Participants were encouraged to follow the Food Guide Pyramid and the equivalent of a National Cholesterol Education Program Step 1 diet, to reduce their weight, and to increase their physical activity. The goals for the participants assigned to the intensive lifestyle intervention were to achieve and maintain a weight reduction of at least 7 percent of initial body weight through a healthy lowcalorie, low-fat diet and to engage in physical activity of moderate intensity, such as brisk walking, for at least 150 minutes per week. A 16-lesson curriculum covering diet, exercise, and behavior modification was designed to help the participants achieve these goals. The curriculum, taught by case managers on a one-to-one basis during the first 24 weeks after enrollment, was flexible, culturally sensitive, and individualized. Subsequent individual sessions (usually monthly) and group sessions with the case managers were designed to reinforce the behavioral changes.

What aspects of the passage indicate that this research followed a randomized controlled experimental study design? In your response, identify the treatment groups, specify which group served as the control, and provide evidence of randomization.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

#### Task 04 💭

<!-- BEGIN QUESTION -->

Next, take a look at the Outcome Measures section (shown below) of the paper (shown below): 

> The primary outcome was diabetes, diagnosed on the basis of an annual oral glucose-tolerance test or a semiannual fasting plasma glucose test, according to the 1997 criteria of the American Diabetes Association: a value for plasma glucose of 126 mg per deciliter (7.0 mmol per liter) or higher in the fasting state or 200 mg per deciliter (11.1 mmol per liter) or higher two hours after a 75-g oral glucose load. In addition to the semiannual measurements, fasting plasma glucose was measured if symptoms suggestive of diabetes developed. The diagnosis required confirmation by a second test, usually within six weeks, according to the same criteria. If diabetes was diagnosed, the participants and their physicians were informed and glucose-tolerance tests were discontinued, but fasting plasma glucose was measured every six months, with glycosylated hemoglobin measured annually. As long as the fasting plasma glucose concentration was less than 140 mg per deciliter, participants were asked to monitor their blood glucose and to continue their assigned study treatment. If the fasting plasma glucose concentration reached or exceeded 140 mg per deciliter, the study medication was discontinued and the participant was referred to his or her physician for treatment. Measurements of glucose and glycosylated hemoglobin ($\text{HbA}_\text{1c}$) were performed centrally. All tests were performed without interrupting the assigned treatment, except that placebo or metformin was not taken on the morning of the test. The investigators and the participants were unaware of the results of these measurements and were informed only if the results exceeded the specified threshold for a change in the treatment. Self-reported levels of leisure physical activity were assessed annually with the Modifiable Activity Questionnaire. The physicalactivity level was calculated as the product of the duration and frequency of each activity (in hours per week), weighted by an estimate of the metabolic equivalent of that activity (MET) and summed for all activities performed, with the result expressed as the average MET-hours per week for the previous year. Usual daily caloric intake during the previous year, including calories from fat, carbohydrate, protein, and other nutrients, was assessed at base line and at one year with the use of a modified version of the Block foodfrequency questionnaire.

What is the primary outcome variable from the study and how was it measured?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Measuring Association

Clinical trials, like the DPP modeled, are randomized controlled experiments. This design is key for establishing causation, but there are more pieces to the research that need to be in place for this argument. Another foundational piece to the story is the measurement of association between the treatments and the outcome(s).

The outcome was measured using the units: number of incident cases per 100 person years. This is an interesting unit of measurement. The number of incident cases is more straightforward, but what does 100 person years mean? From the paper [_Need for and Practical Interpretations of the Person-Year Construct in Neuropsychiatric Research_](https://pmc.ncbi.nlm.nih.gov/articles/PMC6875836/), 

> Person-years refer to the number of years for which persons contribute data. Thus, somebody who was followed up for 5 years would be a five person-year subject. Person-years are totaled in each group to yield the sample size in that group, as expressed in person-years. \[...\] Because the incidence values would be very small for one person-year of follow-up, it is easier to express the values based on the uniform unit of, say, 1000 person-years of follow-up.

So, the number of incident cases per 100 person-years refers to a way of standardizing the occurrence of new cases of a diabetes over time in the population, accounting for both the number of people involved and the length of time they were observed. A person-year is a unit that represents one person being followed for one year, so if someone is observed for five years, they contribute five person-years to the data. By using "per 100 person-years," the researchers can express the incidence rate in a standardized way, making it easier to compare across different groups or studies, particularly when different individuals contribute varying amounts of time. This method allows for clearer interpretations of how often new cases occur in relation to the amount of time people are observed, especially when the event of interest is rare.

#### Task 05 💻

According to the Results section of the paper, the crude incidence was 11.0, 7.8, and 4.8 cases per 100 person-years for the placebo, metformin, and lifestyle intervention groups, respectively. Assign these values to the variable names `placebo`, `metformin`, and `lifestyle`, respectively.

```python
placebo = ...
metformin = ...
lifestyle = ...
```

```python
grader.check("task_05")
```

---

#### Task 06 💭

<!-- BEGIN QUESTION -->

You should notice that all of three of these values are different. How does that observation show evidence of association between the treatment variables and the outcome variable? In terms of the measurements, what would it mean to have no association between the variables?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

#### Task 07 💭

<!-- BEGIN QUESTION -->

Which treatment shows the strongest association with the outcome variable? What does that say about the other treatments?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

### Relative Risk Reduction

With multiple treatments it can be helpful to describe the reduction in the outcome variable relative to the control treatment. Relative Risk Reduction (RRR) is a measure used to describe how much a treatment reduces the risk of an event (such as developing diabetes) in the treatment group compared to the control group. You can calculate RRR as a percentage using the formula

$$\text{Percent Reduction} = \frac{\text{Incidence Rate (Placebo)} − \text{Incidence Rate (Treatment)}}{\text{Incidence Rate (Placebo)}} \times 100$$

#### Task 08 💻

Use the variable names `placebo`, `metformin`, and `lifestyle` to calculate the relative risk reduction for each treatment as a percentage. Assign those percentages to the names `placebo_metformin_RRR` and `placebo_lifestyle_RRR`.

```python
placebo_metformin_RRR = (placebo - metformin) / placebo * 100
placebo_lifestyle_RRR = (placebo - lifestyle) / placebo * 100

print(f"Placebo to Metformin Relative Risk Reduction: {placebo_metformin_RRR:.2f}%")
print(f"Placebo to Lifestyle Relative Risk Reduction: {placebo_lifestyle_RRR:.2f}%")
```

```python
grader.check("task_08")
```

---

#### Task 09 💭

<!-- BEGIN QUESTION -->

Which relative risk reduction was the largest. What does that tell you about that particular treatment?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

### Limiations

The DPP program followed a standard design for establishing a causal relationship. Later in the course you'll learn about statistical significance, and you'll see that they even established significant results which allows them to say that two treatments reduce the risk of Type 2 diabetes. With that said, there is something you need to think about early on in this course. Researchers often are using a limited amount of data that hopefully represents some population of interest. So, we've told you that the problem creates a great argument for the two non-control treatments causing a reduction in the risk of Type 2 diabetes, but who does that conclusion apply to? One way to answer that is to examine the group of people studied.

The following table captures the information in Table 1 from  [_Reduction in the Incidence of Type 2 Diabetes with Lifestyle Intervention or Metformin_](./NEJMoa012512.pdf) 

| CHARACTERISTIC | OVERALL (N=3234) | PLACEBO (N=1082) | METFORMIN (N=1073) | LIFESTYLE (N=1079) |
|---|---|---|---|---|
| **Sex - no. (%)** | | | | |
| Male | 1043 (32.3) | 335 (31.0) | 363 (33.8) | 345 (32.0) |
| Female | 2191 (67.7) | 747 (69.0) | 710 (66.2) | 734 (68.0) |
| **Race or ethnic group - no. (%)** | | | | |
| White | 1768 (54.7) | 586 (54.2) | 602 (56.1) | 580 (53.8) |
| African American | 645 (19.9) | 220 (20.3) | 221 (20.6) | 204 (18.9) |
| Hispanic | 508 (15.7) | 168 (15.5) | 162 (15.1) | 178 (16.5) |
| American Indian | 171 (5.3) | 59 (5.5) | 52 (4.8) | 60 (5.6) |
| Asian† | 142 (4.4) | 49 (4.5) | 36 (3.4) | 57 (5.3) |
| Family history of diabetes - no. (%) | 2243 (69.4) | 758 (70.1) | 733 (68.3) | 752 (69.8)‡ |
| History of gestational diabetes - no. of women (%) | 353 (16.1) | 122 (16.3) | 111 (15.7)‡ | 120 (16.3) |
| **Age - yr (Mean)** | 50.6 ± 10.7 | 50.3 ± 10.4 | 50.9 ± 10.3 | 50.6 ± 11.3 |
| **Weight - kg** | 94.2 ± 20.3 | 94.3 ± 20.2 | 94.3 ± 19.9 | 94.1 ± 20.8 |
| **Body-mass index** | 34.0 ± 6.7 | 34.2 ± 6.7 | 33.9 ± 6.6 | 33.9 ± 6.8 |
| **Waist circumference - cm** | 105.1 ± 14.5 | 105.2 ± 14.3 | 104.9 ± 14.4 | 105.1 ± 14.8 |
| **Waist-to-hip ratio** | 0.92 ± 0.09 | 0.93 ± 0.09 | 0.93 ± 0.09 | 0.92 ± 0.08 |
| **Plasma glucose - mg/dl§** | | | | |
| In the fasting state | 106.5 ± 8.3 | 106.7 ± 8.4 | 106.5 ± 8.5 | 106.3 ± 8.1 |
| Two hours after an oral glucose load | 164.6 ± 17.0 | 164.5 ± 17.1 | 165.1 ± 17.2 | 164.4 ± 16.8 |
| **Glycosylated hemoglobin - %** | 5.91 ± 0.50 | 5.91 ± 0.50 | 5.91 ± 0.50 | 5.91 ± 0.51 |
| **Leisure physical activity - MET-hr/wk¶** | 16.3 ± 25.8 | 17.0 ± 29.0 | 16.4 ± 25.9 | 15.5 ± 22.1 |

_Notes: Plus-minus values are means ±SD. †Twenty Pacific Islanders were included in this category. ‡Information was not available for one participant. §To convert the values for glucose to millimoles per liter, multiply by 0.05551. ¶Data are based on responses to the Modifiable Activity Questionnaire. 16 MET denotes metabolic equivalent. MET-hours represent the average amount of time engaged in specified physical activities multiplied by the MET value of each activity._

This table contains some information on mean and standard deviation (SD). These are values that summarize numerical data. The mean represents the center of the data set and the standard deviation describes the variability of the data set relative to the mean. You'll learn more about these measurements later in the course.

#### Task 10 💭

<!-- BEGIN QUESTION -->

If the data from these participants were used to form the conclusions in the study, what limitations do you think there are about generalizing the results to all humans. In other words, do you see any concerns about the representation of humans in how their characteristics might impact the conclusions?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Your response should prompt more research, not a final conclusion. This is the cyclical nature of this type of work. Remember the data science life cycle from Lecture 01:

```mermaid
flowchart LR
    A[Ask a Question] --> B[Get Data]
    B --> C[Analyze the Data]
    C --> A
```

The conclusions drawn from this study, while valuable, should be applied cautiously and with consideration of the specific demographic and health characteristics of the participants.

---

## Reflection

In this Guided Learning Activity, we explored the fundamental principles of experimental study design, focusing on the distinction between observational studies and randomized controlled experiments. You learned how randomization is crucial in experimental designs to minimize bias and establish causal relationships. The Diabetes Prevention Program (DPP) study provided a concrete example of a well-designed randomized controlled trial, demonstrating the effectiveness of both lifestyle intervention and metformin in reducing the incidence of Type 2 diabetes.  We also introduced the concept of relative risk reduction (RRR) to quantify the impact of different treatments. The DPP study showed that lifestyle intervention was particularly effective, resulting in a 56% reduction in the incidence of diabetes compared to the placebo group. We also practiced writing Python code to assign values to variables and perform calculations, which are essential skills for data analysis.  In future learning activities, we will explore more advanced statistical methods for analyzing data from experimental studies.

---

## License

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a>.

<img src="./by-nc-sa.png" width=100px>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_05 (points: None)
```
>>> import numpy as np
>>> bool(np.all([isinstance(group, (int, float)) for group in [placebo, metformin, lifestyle]]))
True
```
failure_message: "❌ placebo, metformin, lifestyle should
be assigned to numerical values."
success_message: "✅ placebo, metformin, lifestyle are 
assigned to numerical values."

```
>>> placebo == 11.0
True
```
failure_message: "❌ placebo does not seem correct."
success_message: "✅ placebo seems correct!"

```
>>> metformin == 7.8
True
```
failure_message: "❌ metformin does not seem correct."
success_message: "✅ metformin seems correct!"

```
>>> lifestyle == 4.8
True
```
failure_message: "❌ lifestyle does not seem correct."
success_message: "✅ lifestyle seems correct!"



Tests for task_08 (points: None)
```
>>> import numpy as np
>>> bool(np.all([isinstance(group, (int, float)) for group in [placebo_metformin_RRR, placebo_lifestyle_RRR]]))
True
```
failure_message: "❌ placebo_metformin_RRR and placebo_lifestyle_RRR
should be assigned to numerical values."
success_message: "✅ placebo_metformin_RRR and placebo_lifestyle_RRR
are assigned to numerical values."

```
>>> import numpy as np
>>> bool(np.isclose(placebo_metformin_RRR, 56.36, 2))
True
```
failure_message: "❌ placebo_metformin_RRR does not seem correct."
success_message: "✅ placebo_metformin_RRR seems correct!"

```
>>> import numpy as np
>>> bool(np.isclose(placebo_lifestyle_RRR, 29.09, 2))
True
```
failure_message: "❌ placebo_lifestyle_RRR does not seem correct."
success_message: "✅ placebo_lifestyle_RRR seems correct!"

