```python
# Initialize Otter
import otter
grader = otter.Notebook("gla13.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Guided Learning Activity 13: kNN Classification

This Guided Learning Activity is designed for you to complete alongside a Data Ambassador from the course. You might find that it feels like a combination of the lectures and lab assignment. Whether you are participating live or watching the recording of the live meeting, let the Data Ambassador guide you through the following tasks. There will be moments for you to reflect and explore your own ideas as a way to solidify concepts and skills introduced by your instructor. Keep in mind that this is not a graded assignment for MATH 108 by default. If you have any concerns about participation, reach out to your instructor.

---

## Learning Objectives

1. Describe kNN classification
2. Outline a procedure for creating and evaluating a kNN classifier
3. Build a kNN classifier for a real data set
4. Assess the quality of the classifier

---

## Configure the Notebook

Run the following code cell to set up the notebook.

```python
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

---

## Understanding k-Nearest Neighbors (kNN)

---

### Goal

Use the relationship between numerical data and categorical labels to predict a label using only numerical data

---

### How kNN Works

- Predict labels using the majority vote of k nearest known data points
- We measure distance using Euclidean distance
- The k value is a selected during the training process

---

### What Makes a kNN Classifier?

A training set + an optimal k value + selected features

---

### Evaluating Quality
- A metric: Accuracy -> proportion of correctly labeled points
- Data splitting strategy:
  - Training set: builds the model
  - Validation set: optimizes k value
  - Test set: evaluates final performance (used only once)

---

### Improving Performance

- **Standardization**: Ensures no feature dominates due to scale
- **Feature Selection**: Choose non-redundant features that best distinguish between classes

---

### Pipeline

The following is a process for how to convert a labeled data set into a kNN classifier ready to make predictions

```mermaid
flowchart TD
    A["Raw Data"] --> B["Clean Data"]

    B -->|"e.g. 80%"| C["(Pre) Training Set"]
    B -->|"e.g. 20%"| D["Testing Set"]

    C -->|"e.g. 75%"| E["Training Set"]
    C -->|"e.g. 25%"| F["Validation Set"]

    E --> G["Standardize & Select Features"]
    
    subgraph Transform["Transform Data"]
        H1["Transform Training Set"]
        H2["Transform Validation Set"]
        H3["Transform Testing Set"]
    end

    G --> H1
    G --> H2
    G --> H3

    F --> H2
    D --> H3

    H1 --> I["Evaluate Accuracy for several\n k values on Validation Set"]
    H2 --> I
    I --> J["Pick Optimal k"]

    J --> K["Evaluate Final Accuracy on Testing Set"]
    H3 --> K

```

---

### Task 01 📍🔎

<!-- BEGIN QUESTION -->

Feature selection is an important part of the classification process. The following graphic shows two scatterplots using a [wine quality data set](https://archive.ics.uci.edu/dataset/186/wine+quality) and two different sets of features.

<img src="./feature_selection_wine.png" width=800px alt="Two scatterplots showing various class seperation based on two different feature selections">

Which feature selection seems more useful for classifying wine based on the provided category labels?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Detecting Gallstone Disease

---

### Gallstone Disease

<a href="https://en.wikipedia.org/wiki/Gallstone#/media/File:Gallstones.png" target="_blank"><img src="./gallstones.png" width=400px alt="A gallstone blocking a bile duct"></a>

According to [John Hopkins Medicine](https://www.hopkinsmedicine.org/health/conditions-and-diseases/gallstone-disease):

> Gallstone disease is the most common disorder affecting the biliary system, the body's system of transporting bile. Gallstones are solid, pebble-like masses that form in the gallbladder or the biliary tract (the ducts leading from the liver to the small intestine). They form when the bile hardens and are caused by an excess of cholesterol, bile salts or bilirubin.

---

### Research

A recent paper titled [_Early prediction of gallstone disease with a machine learning-based method from bioimpedance and laboratory data_](https://journals.lww.com/md-journal/fulltext/2024/02230/early_prediction_of_gallstone_disease_with_a.40.aspx) showed that vitamin D, C-reactive protein (CRP) level, total body water, and lean mass are crucial features in predicting gallstones.

---

### Data

The UC Irvine Machine Learning Repository hosts [the dataset from this study](https://archive.ics.uci.edu/dataset/1150/gallstone-1), which contains information on 319 individuals, 161 of whom were diagnosed with gallstone disease.

---

### Task 02 📍

The gallstone data is contained within the file `dataset-uci.csv`. Assign `gallstone` to a Table containing that data.

```python
gallstone = ...
gallstone
```

```python
grader.check("task_02")
```

---

### Task 03 📍

Split `gallstone` into two Tables. One table called `pre_train` should contain a random selection of 80% of the data in `gallstone`, and the other table called `test` should contain the remaining rows of `gallstone`.

```python
np.random.seed(123) # Included for reproducibility
...
display(pre_train)
```

```python
grader.check("task_03")
```

---

### Task 04 📍🔎

<!-- BEGIN QUESTION -->

Next, we'd want to clean the pre-training data by checking for missing values, typos, outliers, etc. To keep this activity focused, visualize the distributions of vitamin D, C-reactive protein level, total body water, and lean mass, and make sure there are no outliers or missing values.

**Note**: There is no right way to do this. You are just exploring for any anomalies in the key features outlined by the research paper.

```python
key_features = ['Vitamin D', 'C-Reactive Protein (CRP)',
                'Total Body Water (TBW)', 'Lean Mass (LM) (%)']
...
```

<!-- END QUESTION -->

---

### Outlier Detection

We've created a function for you that identifies the row indices of a table where the data values in those rows are considered outliers by the standard 1.5 IQR method. Run the following code cell to define that function and test it on an array with two outliers.

```python
def detect_outlier_indices(data):
    """
    Detects outliers in an array using the 1.5 * IQR method.
    
    Parameters:
        data: Input data array
        
    Returns:
        indices (list): List of indices corresponding to outlier values
    """
    q1 = percentile(25, data)
    q3 = percentile(75, data)
    iqr = q3 - q1

    lower_bound = q1 - 1.5 * iqr
    upper_bound = q3 + 1.5 * iqr

    outlier_indices = np.where((data < lower_bound) | (data > upper_bound))[0]
    return outlier_indices.tolist()

# Testing the function on an array with 2 outliers
detect_outlier_indices(make_array(-20, 1, 2, 3, 4, 50))
```

The function should return a list of 2 numbers that correspond to the fact that `-20` and `50` are outliers (by the 1.5 IQR method) for the given data set.

---

### Task 05 📍

You should have noticed some potential outliers in two of the features. Use `detect_outlier_indices` to remove the rows of the pre-training data that correspond to those outliers. Assign the resulting table to `clean_pre_train`.

```python
clean_pre_train = ...
clean_pre_train
```

```python
grader.check("task_05")
```

---

### Task 06 📍

Further split `clean_pre_train` into two Tables. One table called `train` should contain a random selection of 75% of the data in `clean_pre_train`, and the other table called `validation` should contain the remaining rows of `clean_pre_train`.

```python
np.random.seed(1234) # Included for reproducibility
...
display(train)
```

```python
grader.check("task_06")
```

---

### Task 07 📍

Next, you need to standardize the data in the 3 data sets based on the values in the training set. We've set up the code to standardize the key features from the data and then transform the train, test, and validation tables by selecting only the standardized features and the Gallstone Status label. You just need to get the average and standard deviation!

```python
for feature in key_features:
    mean_value_train = ...
    SD_value_train = ...

    for tbl_name, tbl in [('train', train), ('validation', validation), ('test', test)]:
        transformed_feature_values = (tbl.column(feature) - mean_value_train) / SD_value_train
        new_column_name = feature + '_standardized'
        tbl = tbl.with_column(new_column_name, transformed_feature_values)

        if tbl_name == 'train':
            train = tbl
        elif tbl_name == 'validation':
            validation = tbl
        else:
            test = tbl

# Transform train, validation, and test sets
standardized_features = ['Gallstone Status', 'Vitamin D_standardized', 
                         'C-Reactive Protein (CRP)_standardized', 'Total Body Water (TBW)_standardized',
                         'Lean Mass (LM) (%)_standardized']

train = train.select(standardized_features)
validation = validation.select(standardized_features)
test = test.select(standardized_features)
display(train_transformed)
```

```python
grader.check("task_07")
```

---

### Task 08 📍

Complete the following functions created in [the MATH 108 textbook](https://ccsf-math-108.github.io/textbook) that were used to build a kNN classifier function called `classify`.

```python
def distance(point1, point2):
    """Returns the distance between point1 and point2
    where each argument is an array 
    consisting of the coordinates of the point"""
    return ...

def all_distances(training, new_point):
    """Returns an array of distances
    between each point in the training set
    and the new point (which is a row of attributes)"""
    attributes = training.drop(...)
    def distance_from_point(row):
        return distance(np.array(new_point), np.array(row))
    return attributes.apply(distance_from_point)

def table_with_distances(training, new_point):
    """Augments the training table 
    with a column of distances from new_point"""
    return training.with_column('Distance', all_distances(training, new_point))

def closest(training, new_point, k):
    """Returns a table of the k rows of the augmented table
    corresponding to the k smallest distances"""
    with_dists = table_with_distances(training, new_point)
    sorted_by_distance = ...
    topk = ...
    return topk

def majority(topkclasses):
    ones = topkclasses.where(..., are.equal_to(1)).num_rows
    zeros = topkclasses.where(..., are.equal_to(0)).num_rows
    if ones > zeros:
        return 1
    else:
        return 0

def classify(training, new_point, k):
    closestk = ...
    topkclasses = ...
    return ...
```

```python
grader.check("task_08")
```

---

### Task 09 📍

Complete the following functions created in [the MATH 108 textbook](https://ccsf-math-108.github.io/textbook) that were used to define a function called `evaluate_accuracy` that is used to evaluate the accuracy of `classify` based on a training set and a test/validation set.

```python
def count_zero(array):
    """Counts the number of 0's in an array"""
    return len(array) - np.count_nonzero(array)

def count_equal(array1, array2):
    """Takes two numerical arrays of equal length
    and counts the indices where the two are equal"""
    return ...

def evaluate_accuracy(training, testing, k):
    test_attributes = testing.drop('Gallstone Status')
    def classify_testrow(row):
        return classify(training, row, k)
    c = test_attributes.apply(classify_testrow)
    return ...
```

```python
grader.check("task_09")
```

---

### Task 10 📍🔎

<!-- BEGIN QUESTION -->

Using the transformed training and validation sets and several values of `k`, create a visual showing the accuracy values for each `k`. Which `k` value seems optimal to use?

_Type your answer here, replacing this text._

```python
accuracy_array = make_array()
ks = [3, 5, 7, 9, 11, 13, 15, 17, 19, 21]
for k in ks:
    accuracy_array = ...

Table().with_columns(
    'k', ks,
    'Accuracy', accuracy_array)
```

<!-- END QUESTION -->

---

### Task 11 📍

Now that you've fine-tuned your kNN classifier by picking out an optional value of `k`, calculate the accuracy of the classifier using the test set. Assign that value to `classifier_accuracy`.

```python
classifier_accuracy = ...
classifier_accuracy
```

```python
grader.check("task_11")
```

---

### Task 12 📍🔎

<!-- BEGIN QUESTION -->

If you evaluate the accuracy with the test set based on `k=5`, you'll actually find the accuracy is higher! Why is it bad to change your classifier to use `k=5`?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Reflection

In this activity, you reviewed the main ideas behind kNN classification, including a workflow for building a classifier. Additionally, you built a kNN classifier and applied it to a data set on gallstone disease detection. Finally, you evaluated the accuracy of the classifier and reflected on how to handle the test set.

---

## License

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a>.

<img src="./by-nc-sa.png" width=100px>

Tests for task_02 (points: None)
```
>>> isinstance(gallstone, Table)
True
```
failure_message: "❌ gallstone is not assigned to a Table."
success_message: "✅ gallstone is assigned to a Table."

```
>>> gallstone
Gallstone Status | Age  | Gender | Comorbidity | Coronary Artery Disease (CAD) | Hypothyroidism | Hyperlipidemia | Diabetes Mellitus (DM) | Height | Weight | Body Mass Index (BMI) | Total Body Water (TBW) | Extracellular Water (ECW) | Intracellular Water (ICW) | Extracellular Fluid/Total Body Water (ECF/TBW) | Total Body Fat Ratio (TBFR) (%) | Lean Mass (LM) (%) | Body Protein Content (Protein) (%) | Visceral Fat Rating (VFR) | Bone Mass (BM) | Muscle Mass (MM) | Obesity (%) | Total Fat Content (TFC) | Visceral Fat Area (VFA) | Visceral Muscle Area (VMA) (Kg) | Hepatic Fat Accumulation (HFA) | Glucose | Total Cholesterol (TC) | Low Density Lipoprotein (LDL) | High Density Lipoprotein (HDL) | Triglyceride | Aspartat Aminotransferaz (AST) | Alanin Aminotransferaz (ALT) | Alkaline Phosphatase (ALP) | Creatinine | Glomerular Filtration Rate (GFR) | C-Reactive Protein (CRP) | Hemoglobin (HGB) | Vitamin D
0                | 50   | 0      | 0           | 0                             | 0              | 0              | 0                      | 185    | 92.8   | 27.1                  | 52.9                   | 21.2                      | 31.7                      | 40                                             | 19.2                            | 80.84              | 18.88                              | 9                         | 3.7            | 71.4             | 23.4        | 17.8                    | 10.6                    | 39.7                            | 0                              | 102     | 250                    | 175                           | 40                             | 134          | 20                             | 22                           | 87                         | 0.82       | 112.47                           | 0                        | 16               | 33
0                | 47   | 0      | 1           | 0                             | 0              | 0              | 0                      | 176    | 94.5   | 30.5                  | 43.1                   | 19.5                      | 23.6                      | 45                                             | 32.8                            | 67.2               | 16.68                              | 15                        | 3.2            | 60.3             | 38.8        | 31                      | 18.4                    | 32.7                            | 0                              | 94      | 172                    | 108                           | 43                             | 103          | 14                             | 13                           | 46                         | 0.87       | 107.1                            | 0                        | 14.4             | 25
0                | 61   | 0      | 0           | 0                             | 0              | 0              | 0                      | 171    | 91.1   | 31.2                  | 47.2                   | 20.1                      | 27.1                      | 43                                             | 27.3                            | 72.67              | 16.35                              | 15                        | 3.3            | 62.9             | 41.7        | 24.9                    | 16.2                    | 34                              | 0                              | 103     | 179                    | 124                           | 43                             | 69           | 18                             | 14                           | 66                         | 1.25       | 65.51                            | 0                        | 16.2             | 30.2
0                | 41   | 0      | 0           | 0                             | 0              | 0              | 0                      | 168    | 67.7   | 24                    | 41.4                   | 17                        | 24.4                      | 41                                             | 15.8                            | 84.19              | 16.9                               | 6                         | 2.9            | 54.1             | 9           | 10.7                    | 6.5                     | 29.2                            | 1                              | 69      | 173                    | 73                            | 59                             | 53           | 20                             | 12                           | 34                         | 1.02       | 94.1                             | 0                        | 15.4             | 35.4
0                | 42   | 0      | 0           | 0                             | 0              | 0              | 0                      | 178    | 89.6   | 28.3                  | 51.4                   | 20                        | 31.4                      | 39                                             | 20                              | 80.02              | 16.81                              | 8                         | 3.5            | 68.2             | 28.6        | 17.9                    | 10.4                    | 37.4                            | 2                              | 109     | 205                    | 154                           | 30                             | 326          | 27                             | 54                           | 71                         | 0.82       | 112.47                           | 0                        | 16.8             | 40.6
0                | 96   | 0      | 0           | 0                             | 0              | 0              | 0                      | 155    | 49     | 20.4                  | 34                     | 15.7                      | 18.3                      | 46                                             | 6.3                             | 93.67              | 18.48                              | 12                        | 2.4            | 43.5             | 7.4         | 3.1                     | 0.9                     | 25.9                            | 0                              | 74      | 133                    | 91                            | 30                             | 65           | 13                             | 13                           | 60                         | 1.46       | 43.74                            | 0                        | 11               | 45.8
0                | 37   | 0      | 0           | 0                             | 0              | 0              | 0                      | 185    | 67.1   | 19.6                  | 41.8                   | 17.1                      | 24.7                      | 41                                             | 11.3                            | 88.67              | 19.9                               | 3                         | 3              | 56.5             | 10.9        | 7.6                     | 4.6                     | 31.4                            | 0                              | 94      | 166                    | 111                           | 40                             | 60           | 15                             | 14                           | 72                         | 0.77       | 118.25                           | 0                        | 13.8             | 20
0                | 41   | 0      | 0           | 0                             | 0              | 0              | 0                      | 176    | 114    | 36.8                  | 57.6                   | 23.5                      | 34.1                      | 41                                             | 31.8                            | 68.16              | 12.66                              | 16                        | 3.8            | 73.9             | 67.4        | 36.3                    | 23.7                    | 37.5                            | 3                              | 125     | 246                    | 103.6                         | 34                             | 464          | 26                             | 28                           | 69                         | 1.3        | 73.73                            | 0.11                     | 16.5             | 24.5
0                | 38   | 0      | 0           | 0                             | 0              | 0              | 0                      | 186    | 93.7   | 27.1                  | 51.9                   | 20.9                      | 31                        | 40                                             | 21.5                            | 78.55              | 17.42                              | 8                         | 3.6            | 70               | 23.1        | 20.1                    | 12.7                    | 37.8                            | 2                              | 93      | 290                    | 122                           | 33                             | 679          | 68                             | 102                          | 70                         | 0.91       | 110.63                           | 1.57                     | 16.5             | 22.7
0                | 38   | 0      | 0           | 0                             | 0              | 0              | 0                      | 171    | 68.6   | 23.5                  | 39.5                   | 16.6                      | 22.9                      | 42                                             | 19.2                            | 80.76              | 17.28                              | 6                         | 2.8            | 52.6             | 6.7         | 13.2                    | 8.2                     | 28.8                            | 0                              | 93      | 239                    | 169                           | 43                             | 129          | 19                             | 34                           | 75                         | 0.91       | 110.63                           | 0                        | 16.6             | 15.6
... (309 rows omitted)
```
failure_message: "❌ gallstone does not seem correct."
success_message: "✅ gallstone seems correct."



Tests for task_03 (points: None)
```
>>> isinstance(pre_train, Table)
True
```
failure_message: "❌ pre_train is not assigned to a Table."
success_message: "✅ pre_train is assigned to a Table."

```
>>> isinstance(test, Table)
True
```
failure_message: "❌ test is not assigned to a Table."
success_message: "✅ test is assigned to a Table."

```
>>> pre_train.num_rows
255
```
failure_message: "❌ pre-train does not have the correct number of rows."
success_message: "✅ pre-train has the correct number of rows."

```
>>> test.num_rows
64
```
failure_message: "❌ test does not have the correct number of rows."
success_message: "✅ test has the correct number of rows."

```
>>> pre_train
Gallstone Status | Age  | Gender | Comorbidity | Coronary Artery Disease (CAD) | Hypothyroidism | Hyperlipidemia | Diabetes Mellitus (DM) | Height | Weight | Body Mass Index (BMI) | Total Body Water (TBW) | Extracellular Water (ECW) | Intracellular Water (ICW) | Extracellular Fluid/Total Body Water (ECF/TBW) | Total Body Fat Ratio (TBFR) (%) | Lean Mass (LM) (%) | Body Protein Content (Protein) (%) | Visceral Fat Rating (VFR) | Bone Mass (BM) | Muscle Mass (MM) | Obesity (%) | Total Fat Content (TFC) | Visceral Fat Area (VFA) | Visceral Muscle Area (VMA) (Kg) | Hepatic Fat Accumulation (HFA) | Glucose | Total Cholesterol (TC) | Low Density Lipoprotein (LDL) | High Density Lipoprotein (HDL) | Triglyceride | Aspartat Aminotransferaz (AST) | Alanin Aminotransferaz (ALT) | Alkaline Phosphatase (ALP) | Creatinine | Glomerular Filtration Rate (GFR) | C-Reactive Protein (CRP) | Hemoglobin (HGB) | Vitamin D
0                | 50   | 0      | 1           | 0                             | 0              | 0              | 1                      | 178    | 110.7  | 34.9                  | 53.9                   | 22.7                      | 31.2                      | 42                                             | 31.9                            | 68.11              | 15.2                               | 17                        | 3.7            | 71.7             | 58.8        | 35.3                    | 23                      | 37.2                            | 2                              | 374     | 225                    | 128                           | 33                             | 692          | 24                             | 24                           | 97                         | 0.83       | 107.28                           | 0.1                      | 17.9             | 11.1
0                | 43   | 0      | 0           | 0                             | 0              | 0              | 0                      | 168    | 73.4   | 26                    | 45.5                   | 18.1                      | 27.4                      | 40                                             | 15.4                            | 84.6               | 16.44                              | 6                         | 3.1            | 59               | 18.2        | 11.3                    | 6.6                     | 31.9                            | 0                              | 92      | 259                    | 115                           | 50                             | 206          | 16                             | 30                           | 103                        | 0.84       | 110.97                           | 0                        | 16.6             | 22.4
0                | 49   | 1      | 1           | 0                             | 1              | 0              | 0                      | 158    | 92.5   | 37.1                  | 38.1                   | 18                        | 20.1                      | 47                                             | 42.2                            | 57.84              | 12.43                              | 11                        | 2.7            | 50.8             | 68.5        | 39                      | 17.3                    | 29.8                            | 1                              | 92      | 208                    | 129                           | 57                             | 82           | 18                             | 18                           | 92                         | 0.63       | 109.77                           | 0.78                     | 10.7             | 21.8
0                | 47   | 1      | 1           | 0                             | 0              | 0              | 0                      | 169    | 79     | 27.7                  | 35.7                   | 16.3                      | 19.4                      | 46                                             | 36.6                            | 63.42              | 13.6                               | 7                         | 2.5            | 47.6             | 25.8        | 28.9                    | 14.5                    | 27.5                            | 1                              | 83      | 172                    | 100                           | 54                             | 116          | 9                              | 8                            | 119                        | 0.6        | 11.34                            | 0.31                     | 14               | 10.9667
1                | 35   | 1      | 0           | 0                             | 0              | 0              | 0                      | 165    | 97.7   | 35.9                  | 37.5                   | 15                        | 23                        | 40                                             | 45.04                           | 54.96              | 12.57                              | 10                        | 2.2            | 51               | 52.7        | 44                      | 23.1                    | 29.8                            | 2                              | 85      | 187                    | 102                           | 77                             | 73           | 11                             | 8                            | 71                         | 0.57       | 121.4                            | 21.4                     | 12.6             | 31.5
0                | 55   | 0      | 0           | 0                             | 0              | 0              | 0                      | 176    | 90.7   | 29.3                  | 48.9                   | 20.3                      | 28.6                      | 42                                             | 23.9                            | 76.07              | 17.44                              | 12                        | 3.4            | 65.6             | 33.2        | 21.7                    | 13.2                    | 36.6                            | 2                              | 176     | 172                    | 98                            | 30                             | 330          | 29                             | 57                           | 64                         | 0.81       | 105.43                           | 0.23                     | 14.3             | 33.8
1                | 51   | 0      | 1           | 0                             | 0              | 0              | 1                      | 165    | 68.2   | 25.1                  | 39.7                   | 16.9                      | 22.8                      | 43                                             | 19.1                            | 80.94              | 17.1                               | 9                         | 2.8            | 52.4             | 13.9        | 13                      | 8                       | 32.5568                         | 0                              | 168     | 143                    | 94                            | 32                             | 91           | 14                             | 14                           | 86                         | 0.73       | 89.0013                          | 1.25                     | 13.9             | 35.4
0                | 48   | 0      | 0           | 0                             | 0              | 0              | 0                      | 168    | 86.8   | 30.8                  | 44.4                   | 19                        | 25.4                      | 43                                             | 28.1                            | 71.89              | 15.49                              | 13                        | 3.1            | 59.3             | 39.8        | 24.4                    | 14.5                    | 32.3                            | 0                              | 91      | 206                    | 152                           | 38                             | 115          | 20                             | 26                           | 73                         | 1.03       | 89.6                             | 0                        | 15.9             | 24
1                | 40   | 0      | 1           | 0                             | 0              | 0              | 1                      | 182    | 110.6  | 33.4                  | 51.8                   | 22.3                      | 29.5                      | 43                                             | 32.8                            | 67.18              | 15.4                               | 15                        | 3.7            | 74.3             | 51.7        | 36.3                    | 21.8                    | 31.9838                         | 2                              | 91      | 180                    | 122                           | 42                             | 145          | 17                             | 36                           | 49                         | 0.86       | 97.924                           | 3.7                      | 15.7             | 13
1                | 57   | 0      | 1           | 0                             | 0              | 0              | 1                      | 168    | 93.4   | 33.1                  | 48.7                   | 20                        | 29                        | 41.06                                          | 27.09                           | 72.81              | 16.15                              | 15                        | 3.4            | 64.7             | 29.76       | 25.3                    | 14.96                   | 34.8                            | 2                              | 128     | 88                     | 40                            | 44                             | 55           | 18                             | 36                           | 73                         | 0.61       | 11                               | 2                        | 16               | 20.2
... (245 rows omitted)
```
failure_message: "❌ pre_train does not seem correct based on this random seed value."
success_message: "✅ pre_train seems correct based on this random seed value."

```
>>> test
Gallstone Status | Age  | Gender | Comorbidity | Coronary Artery Disease (CAD) | Hypothyroidism | Hyperlipidemia | Diabetes Mellitus (DM) | Height | Weight | Body Mass Index (BMI) | Total Body Water (TBW) | Extracellular Water (ECW) | Intracellular Water (ICW) | Extracellular Fluid/Total Body Water (ECF/TBW) | Total Body Fat Ratio (TBFR) (%) | Lean Mass (LM) (%) | Body Protein Content (Protein) (%) | Visceral Fat Rating (VFR) | Bone Mass (BM) | Muscle Mass (MM) | Obesity (%) | Total Fat Content (TFC) | Visceral Fat Area (VFA) | Visceral Muscle Area (VMA) (Kg) | Hepatic Fat Accumulation (HFA) | Glucose | Total Cholesterol (TC) | Low Density Lipoprotein (LDL) | High Density Lipoprotein (HDL) | Triglyceride | Aspartat Aminotransferaz (AST) | Alanin Aminotransferaz (ALT) | Alkaline Phosphatase (ALP) | Creatinine | Glomerular Filtration Rate (GFR) | C-Reactive Protein (CRP) | Hemoglobin (HGB) | Vitamin D
1                | 54   | 1      | 0           | 0                             | 0              | 0              | 0                      | 155    | 80.2   | 33.4                  | 32.8                   | 15                        | 18                        | 45.73                                          | 46.13                           | 53.87              | 9.63                               | 12                        | 2.2            | 41               | 36.23       | 37                      | 19.78                   | 23.7                            | 2                              | 111     | 138                    | 93                            | 31                             | 104          | 37                             | 40                           | 39                         | 0.89       | 77                               | 6.5                      | 14.2             | 10.4
0                | 32   | 0      | 0           | 0                             | 0              | 0              | 0                      | 170    | 64.3   | 22.2                  | 39.8                   | 16.5                      | 23.3                      | 41                                             | 15.3                            | 84.76              | 16.67                              | 4                         | 2.7            | 51.8             | 1.1         | 9.8                     | 5.9                     | 27.6                            | 1                              | 91      | 185                    | 125                           | 35                             | 165          | 23                             | 22                           | 69                         | 0.67       | 127.22                           | 0.5                      | 16               | 22.925
0                | 57   | 1      | 1           | 0                             | 1              | 0              | 1                      | 158    | 79.8   | 32                    | 34.1                   | 16.1                      | 18                        | 47                                             | 39.6                            | 60.4               | 13.92                              | 10                        | 2.4            | 45.8             | 45.4        | 31.6                    | 14.7                    | 26.3                            | 2                              | 168     | 272                    | 179                           | 57                             | 302          | 35                             | 27                           | 80                         | 0.71       | 98.5                             | 0.54                     | 12               | 24.6
0                | 53   | 1      | 0           | 0                             | 0              | 0              | 0                      | 160    | 78.4   | 30.6                  | 33.9                   | 15.9                      | 18                        | 47                                             | 39                              | 60.97              | 13.95                              | 9                         | 2.4            | 45.4             | 39.3        | 30.6                    | 14.5                    | 26.2                            | 0                              | 119     | 211                    | 139                           | 50                             | 117          | 21                             | 32                           | 131                        | 0.7        | 103.99                           | 0                        | 15               | 7.8
0                | 52   | 0      | 0           | 0                             | 0              | 0              | 0                      | 172    | 96.7   | 32.7                  | 50.7                   | 20.8                      | 29.9                      | 41                                             | 27.3                            | 72.7               | 15.76                              | 14                        | 3.5            | 66.8             | 48.5        | 26.4                    | 15.8                    | 36.4                            | 2                              | 107     | 174                    | 110                           | 33                             | 245          | 44                             | 54                           | 76                         | 0.81       | 106.09                           | 0                        | 15.2             | 23.9
0                | 41   | 0      | 0           | 0                             | 0              | 0              | 0                      | 168    | 67.7   | 24                    | 41.4                   | 17                        | 24.4                      | 41                                             | 15.8                            | 84.19              | 16.9                               | 6                         | 2.9            | 54.1             | 9           | 10.7                    | 6.5                     | 29.2                            | 1                              | 69      | 173                    | 73                            | 59                             | 53           | 20                             | 12                           | 34                         | 1.02       | 94.1                             | 0                        | 15.4             | 35.4
0                | 38   | 1      | 1           | 0                             | 0              | 0              | 1                      | 157    | 54.8   | 22.2                  | 29.3                   | 12.3                      | 17                        | 42                                             | 25.4                            | 74.64              | 15.72                              | 3                         | 2.1            | 38.8             | 1.1         | 13.9                    | 5.5                     | 22.7                            | 1                              | 92      | 259                    | 204                           | 37                             | 135          | 12                             | 14                           | 47                         | 0.96       | 113.2                            | 0                        | 14.1             | 32
1                | 59   | 1      | 1           | 0                             | 0              | 0              | 0                      | 150    | 53.6   | 23.8                  | 27                     | 12.2                      | 14.8                      | 45                                             | 28.7                            | 71.27              | 16.48                              | 6                         | 2              | 36.2             | 8.3         | 15.4                    | 6.4                     | 21                              | 0                              | 91      | 240                    | 150                           | 73                             | 95           | 14                             | 13                           | 79                         | 0.65       | 101.3                            | 1.9                      | 13.6             | 8.7
0                | 24   | 0      | 0           | 0                             | 0              | 0              | 0                      | 174    | 68.4   | 22.6                  | 41.8                   | 16.8                      | 25                        | 40                                             | 15.9                            | 84.06              | 16.82                              | 3                         | 2.9            | 54.6             | 2.7         | 10.9                    | 7                       | 28.6                            | 0                              | 85      | 172                    | 111                           | 43                             | 71           | 16                             | 23                           | 59                         | 0.96       | 113.2                            | 0                        | 15.4             | 35
0                | 47   | 0      | 1           | 0                             | 0              | 0              | 0                      | 176    | 94.5   | 30.5                  | 43.1                   | 19.5                      | 23.6                      | 45                                             | 32.8                            | 67.2               | 16.68                              | 15                        | 3.2            | 60.3             | 38.8        | 31                      | 18.4                    | 32.7                            | 0                              | 94      | 172                    | 108                           | 43                             | 103          | 14                             | 13                           | 46                         | 0.87       | 107.1                            | 0                        | 14.4             | 25
... (54 rows omitted)
```
failure_message: "❌ test does not seem correct based on this random seed value."
success_message: "✅ test seems correct based on this random seed value."



Tests for task_05 (points: None)
```
>>> isinstance(clean_pre_train, Table)
True
```
failure_message: "❌ clean_pre_train is not assigned to a Table."
success_message: "✅ clean_pre_train is assigned to a Table."

```
>>> set(clean_pre_train.labels) == set(pre_train.labels)
True
```
failure_message: "❌ clean_pre_train has the same labels as pre_train."
success_message: "✅ clean_pre_train has the same labels as pre_train."

```
>>> clean_pre_train
Gallstone Status | Age  | Gender | Comorbidity | Coronary Artery Disease (CAD) | Hypothyroidism | Hyperlipidemia | Diabetes Mellitus (DM) | Height | Weight | Body Mass Index (BMI) | Total Body Water (TBW) | Extracellular Water (ECW) | Intracellular Water (ICW) | Extracellular Fluid/Total Body Water (ECF/TBW) | Total Body Fat Ratio (TBFR) (%) | Lean Mass (LM) (%) | Body Protein Content (Protein) (%) | Visceral Fat Rating (VFR) | Bone Mass (BM) | Muscle Mass (MM) | Obesity (%) | Total Fat Content (TFC) | Visceral Fat Area (VFA) | Visceral Muscle Area (VMA) (Kg) | Hepatic Fat Accumulation (HFA) | Glucose | Total Cholesterol (TC) | Low Density Lipoprotein (LDL) | High Density Lipoprotein (HDL) | Triglyceride | Aspartat Aminotransferaz (AST) | Alanin Aminotransferaz (ALT) | Alkaline Phosphatase (ALP) | Creatinine | Glomerular Filtration Rate (GFR) | C-Reactive Protein (CRP) | Hemoglobin (HGB) | Vitamin D
0                | 50   | 0      | 1           | 0                             | 0              | 0              | 1                      | 178    | 110.7  | 34.9                  | 53.9                   | 22.7                      | 31.2                      | 42                                             | 31.9                            | 68.11              | 15.2                               | 17                        | 3.7            | 71.7             | 58.8        | 35.3                    | 23                      | 37.2                            | 2                              | 374     | 225                    | 128                           | 33                             | 692          | 24                             | 24                           | 97                         | 0.83       | 107.28                           | 0.1                      | 17.9             | 11.1
0                | 43   | 0      | 0           | 0                             | 0              | 0              | 0                      | 168    | 73.4   | 26                    | 45.5                   | 18.1                      | 27.4                      | 40                                             | 15.4                            | 84.6               | 16.44                              | 6                         | 3.1            | 59               | 18.2        | 11.3                    | 6.6                     | 31.9                            | 0                              | 92      | 259                    | 115                           | 50                             | 206          | 16                             | 30                           | 103                        | 0.84       | 110.97                           | 0                        | 16.6             | 22.4
0                | 49   | 1      | 1           | 0                             | 1              | 0              | 0                      | 158    | 92.5   | 37.1                  | 38.1                   | 18                        | 20.1                      | 47                                             | 42.2                            | 57.84              | 12.43                              | 11                        | 2.7            | 50.8             | 68.5        | 39                      | 17.3                    | 29.8                            | 1                              | 92      | 208                    | 129                           | 57                             | 82           | 18                             | 18                           | 92                         | 0.63       | 109.77                           | 0.78                     | 10.7             | 21.8
0                | 47   | 1      | 1           | 0                             | 0              | 0              | 0                      | 169    | 79     | 27.7                  | 35.7                   | 16.3                      | 19.4                      | 46                                             | 36.6                            | 63.42              | 13.6                               | 7                         | 2.5            | 47.6             | 25.8        | 28.9                    | 14.5                    | 27.5                            | 1                              | 83      | 172                    | 100                           | 54                             | 116          | 9                              | 8                            | 119                        | 0.6        | 11.34                            | 0.31                     | 14               | 10.9667
0                | 55   | 0      | 0           | 0                             | 0              | 0              | 0                      | 176    | 90.7   | 29.3                  | 48.9                   | 20.3                      | 28.6                      | 42                                             | 23.9                            | 76.07              | 17.44                              | 12                        | 3.4            | 65.6             | 33.2        | 21.7                    | 13.2                    | 36.6                            | 2                              | 176     | 172                    | 98                            | 30                             | 330          | 29                             | 57                           | 64                         | 0.81       | 105.43                           | 0.23                     | 14.3             | 33.8
1                | 51   | 0      | 1           | 0                             | 0              | 0              | 1                      | 165    | 68.2   | 25.1                  | 39.7                   | 16.9                      | 22.8                      | 43                                             | 19.1                            | 80.94              | 17.1                               | 9                         | 2.8            | 52.4             | 13.9        | 13                      | 8                       | 32.5568                         | 0                              | 168     | 143                    | 94                            | 32                             | 91           | 14                             | 14                           | 86                         | 0.73       | 89.0013                          | 1.25                     | 13.9             | 35.4
0                | 48   | 0      | 0           | 0                             | 0              | 0              | 0                      | 168    | 86.8   | 30.8                  | 44.4                   | 19                        | 25.4                      | 43                                             | 28.1                            | 71.89              | 15.49                              | 13                        | 3.1            | 59.3             | 39.8        | 24.4                    | 14.5                    | 32.3                            | 0                              | 91      | 206                    | 152                           | 38                             | 115          | 20                             | 26                           | 73                         | 1.03       | 89.6                             | 0                        | 15.9             | 24
1                | 40   | 0      | 1           | 0                             | 0              | 0              | 1                      | 182    | 110.6  | 33.4                  | 51.8                   | 22.3                      | 29.5                      | 43                                             | 32.8                            | 67.18              | 15.4                               | 15                        | 3.7            | 74.3             | 51.7        | 36.3                    | 21.8                    | 31.9838                         | 2                              | 91      | 180                    | 122                           | 42                             | 145          | 17                             | 36                           | 49                         | 0.86       | 97.924                           | 3.7                      | 15.7             | 13
1                | 57   | 0      | 1           | 0                             | 0              | 0              | 1                      | 168    | 93.4   | 33.1                  | 48.7                   | 20                        | 29                        | 41.06                                          | 27.09                           | 72.81              | 16.15                              | 15                        | 3.4            | 64.7             | 29.76       | 25.3                    | 14.96                   | 34.8                            | 2                              | 128     | 88                     | 40                            | 44                             | 55           | 18                             | 36                           | 73                         | 0.61       | 11                               | 2                        | 16               | 20.2
1                | 57   | 1      | 1           | 0                             | 0              | 1              | 0                      | 162    | 74.7   | 28.5                  | 32.7                   | 15.5                      | 17.2                      | 47                                             | 38.3                            | 61.7               | 14.11                              | 9                         | 2.3            | 43.8             | 29.5        | 28.6                    | 14.3                    | 28.6892                         | 3                              | 95      | 305                    | 224                           | 47                             | 216          | 16                             | 16                           | 82                         | 0.64       | 101.8                            | 2.9                      | 12.7             | 17.6
... (217 rows omitted)
```
failure_message: "❌ clean_pre_train does not seem correct based on this random seed."
success_message: "✅ clean_pre_train seems correct based on this random seed."



Tests for task_06 (points: None)
```
>>> isinstance(train, Table)
True
```
failure_message: "❌ train is not assigned to a Table."
success_message: "✅ train is assigned to a Table."

```
>>> isinstance(validation, Table)
True
```
failure_message: "❌ validation is not assigned to a Table."
success_message: "✅ validation is assigned to a Table."

```
>>> train.num_rows
170
```
failure_message: "❌ train does not have the correct number of rows."
success_message: "✅ train has the correct number of rows."

```
>>> validation.num_rows
57
```
failure_message: "❌ validation does not have the correct number of rows."
success_message: "✅ validation has the correct number of rows."

```
>>> train
Gallstone Status | Age  | Gender | Comorbidity | Coronary Artery Disease (CAD) | Hypothyroidism | Hyperlipidemia | Diabetes Mellitus (DM) | Height | Weight | Body Mass Index (BMI) | Total Body Water (TBW) | Extracellular Water (ECW) | Intracellular Water (ICW) | Extracellular Fluid/Total Body Water (ECF/TBW) | Total Body Fat Ratio (TBFR) (%) | Lean Mass (LM) (%) | Body Protein Content (Protein) (%) | Visceral Fat Rating (VFR) | Bone Mass (BM) | Muscle Mass (MM) | Obesity (%) | Total Fat Content (TFC) | Visceral Fat Area (VFA) | Visceral Muscle Area (VMA) (Kg) | Hepatic Fat Accumulation (HFA) | Glucose | Total Cholesterol (TC) | Low Density Lipoprotein (LDL) | High Density Lipoprotein (HDL) | Triglyceride | Aspartat Aminotransferaz (AST) | Alanin Aminotransferaz (ALT) | Alkaline Phosphatase (ALP) | Creatinine | Glomerular Filtration Rate (GFR) | C-Reactive Protein (CRP) | Hemoglobin (HGB) | Vitamin D
0                | 35   | 1      | 0           | 0                             | 0              | 0              | 0                      | 164    | 79.3   | 29.5                  | 36.7                   | 16.3                      | 20.4                      | 44                                             | 35.4                            | 64.56              | 13.57                              | 6                         | 2.6            | 48.6             | 34          | 28.1                    | 12.9                    | 28.6                            | 1                              | 91      | 227                    | 160                           | 41                             | 129          | 16                             | 20                           | 61                         | 0.88       | 87.83                            | 0                        | 16               | 7.7
0                | 48   | 0      | 0           | 0                             | 0              | 0              | 0                      | 168    | 86.8   | 30.8                  | 44.4                   | 19                        | 25.4                      | 43                                             | 28.1                            | 71.89              | 15.49                              | 13                        | 3.1            | 59.3             | 39.8        | 24.4                    | 14.5                    | 32.3                            | 0                              | 91      | 206                    | 152                           | 38                             | 115          | 20                             | 26                           | 73                         | 1.03       | 89.6                             | 0                        | 15.9             | 24
1                | 48   | 1      | 1           | 0                             | 0              | 0              | 1                      | 154    | 76.8   | 32.4                  | 37.1                   | 16.5                      | 20.6                      | 44                                             | 32                              | 67.97              | 14.7                               | 8                         | 2.6            | 49.6             | 47.1        | 24.6                    | 9.1                     | 28.4                            | 2                              | 104     | 143                    | 84                            | 36                             | 204          | 19                             | 27                           | 74                         | 0.56       | 113                              | 1.1                      | 13.3             | 7
1                | 40   | 0      | 1           | 0                             | 0              | 0              | 1                      | 182    | 110.6  | 33.4                  | 51.8                   | 22.3                      | 29.5                      | 43                                             | 32.8                            | 67.18              | 15.4                               | 15                        | 3.7            | 74.3             | 51.7        | 36.3                    | 21.8                    | 31.9838                         | 2                              | 91      | 180                    | 122                           | 42                             | 145          | 17                             | 36                           | 49                         | 0.86       | 97.924                           | 3.7                      | 15.7             | 13
0                | 50   | 1      | 0           | 0                             | 0              | 0              | 0                      | 163    | 74.8   | 28.2                  | 36.3                   | 16.1                      | 20.1                      | 44                                             | 31.7                            | 68.32              | 15.55                              | 7                         | 2.6            | 48.5             | 27.9        | 23.7                    | 10.2                    | 28.5                            | 0                              | 118     | 171                    | 122                           | 40                             | 98           | 16                             | 16                           | 63                         | 0.53       | 111.9                            | 1.71                     | 14.1             | 31.85
0                | 58   | 0      | 1           | 0                             | 0              | 0              | 1                      | 170    | 79     | 27.3                  | 41.8                   | 18.1                      | 23.7                      | 43                                             | 24.3                            | 75.7               | 18.09                              | 12                        | 3              | 56.8             | 24.2        | 19.2                    | 11.4                    | 32.2                            | 0                              | 115     | 235                    | 161                           | 47                             | 104          | 27                             | 15                           | 15                         | 1          | 85.09                            | 1.3                      | 15.5             | 35
0                | 50   | 0      | 0           | 0                             | 0              | 0              | 0                      | 172    | 97.3   | 32.9                  | 46.5                   | 20.3                      | 26.2                      | 44                                             | 31.7                            | 68.35              | 16.32                              | 16                        | 3.3            | 63.2             | 49.5        | 30.8                    | 18.3                    | 34.2                            | 2                              | 126     | 194                    | 145                           | 26                             | 132          | 15                             | 29                           | 50                         | 0.83       | 106.6                            | 0.55                     | 15.6             | 24.4286
1                | 56   | 0      | 0           | 0                             | 0              | 0              | 0                      | 171    | 73.1   | 25                    | 42.4                   | 18                        | 24.4                      | 42                                             | 18.4                            | 81.53              | 18.47                              | 9                         | 3              | 56.6             | 13.7        | 13.5                    | 8.2                     | 31.5                            | 0                              | 107     | 250                    | 166                           | 64                             | 93           | 13                             | 18                           | 46                         | 1.2        | 86.26                            | 0                        | 14.6             | 19.4
1                | 51   | 0      | 1           | 0                             | 0              | 0              | 1                      | 165    | 68.2   | 25.1                  | 39.7                   | 16.9                      | 22.8                      | 43                                             | 19.1                            | 80.94              | 17.1                               | 9                         | 2.8            | 52.4             | 13.9        | 13                      | 8                       | 32.5568                         | 0                              | 168     | 143                    | 94                            | 32                             | 91           | 14                             | 14                           | 86                         | 0.73       | 89.0013                          | 1.25                     | 13.9             | 35.4
0                | 50   | 0      | 0           | 0                             | 0              | 0              | 0                      | 174    | 85.5   | 28.2                  | 47.3                   | 19.4                      | 27.9                      | 41                                             | 22.4                            | 77.54              | 17.41                              | 11                        | 3.3            | 63               | 28.4        | 19.2                    | 11.6                    | 34.9                            | 2                              | 102     | 226                    | 147                           | 61                             | 78           | 20                             | 34                           | 83                         | 0.97       | 95.1                             | 0.32                     | 16               | 29.25
... (160 rows omitted)
```
failure_message: "❌ train does not seem correct based on this random seed value."
success_message: "✅ train seems correct based on this random seed value."

```
>>> validation
Gallstone Status | Age  | Gender | Comorbidity | Coronary Artery Disease (CAD) | Hypothyroidism | Hyperlipidemia | Diabetes Mellitus (DM) | Height | Weight | Body Mass Index (BMI) | Total Body Water (TBW) | Extracellular Water (ECW) | Intracellular Water (ICW) | Extracellular Fluid/Total Body Water (ECF/TBW) | Total Body Fat Ratio (TBFR) (%) | Lean Mass (LM) (%) | Body Protein Content (Protein) (%) | Visceral Fat Rating (VFR) | Bone Mass (BM) | Muscle Mass (MM) | Obesity (%) | Total Fat Content (TFC) | Visceral Fat Area (VFA) | Visceral Muscle Area (VMA) (Kg) | Hepatic Fat Accumulation (HFA) | Glucose | Total Cholesterol (TC) | Low Density Lipoprotein (LDL) | High Density Lipoprotein (HDL) | Triglyceride | Aspartat Aminotransferaz (AST) | Alanin Aminotransferaz (ALT) | Alkaline Phosphatase (ALP) | Creatinine | Glomerular Filtration Rate (GFR) | C-Reactive Protein (CRP) | Hemoglobin (HGB) | Vitamin D
0                | 71   | 0      | 1           | 0                             | 0              | 0              | 0                      | 172    | 78     | 26.4                  | 49.5                   | 19.7                      | 29.8                      | 40                                             | 12.1                            | 87.95              | 19.03                              | 10                        | 3.4            | 65.2             | 19.8        | 9.4                     | 4.4                     | 37.4                            | 2                              | 110     | 253                    | 182                           | 47                             | 190          | 25                             | 26                           | 26                         | 1.01       | 79.51                            | 0                        | 14.4             | 31.06
1                | 31   | 1      | 0           | 0                             | 0              | 0              | 0                      | 157    | 53.4   | 21.7                  | 29.3                   | 12.3                      | 17                        | 42                                             | 23.6                            | 76.4               | 15.96                              | 2                         | 2.1            | 38.7             | 1.5         | 12.6                    | 5.1                     | 22.2                            | 0                              | 96      | 176                    | 99                            | 58                             | 64           | 24                             | 16                           | 38                         | 0.5        | 128.5                            | 0                        | 12.5             | 24
0                | 43   | 0      | 0           | 0                             | 0              | 0              | 0                      | 176    | 75.7   | 24.4                  | 39.4                   | 17.1                      | 22.3                      | 43                                             | 24.3                            | 75.69              | 18.12                              | 9                         | 2.9            | 54.4             | 11.2        | 18.4                    | 10.4                    | 30.8                            | 0                              | 99      | 167                    | 110                           | 40                             | 63           | 19                             | 18                           | 76                         | 0.98       | 98.12                            | 0.05                     | 15               | 28
1                | 52   | 1      | 0           | 0                             | 0              | 0              | 0                      | 150    | 74.7   | 33.2                  | 33.6                   | 15                        | 19                        | 44.64                                          | 38.55                           | 61.45              | 12.65                              | 10                        | 2.3            | 43.5             | 35.5        | 28.8                    | 13.42                   | 25.3                            | 0                              | 100     | 190                    | 118                           | 58                             | 68           | 25                             | 18                           | 66                         | 0.75       | 95.7                             | 0                        | 11.6             | 12.1
0                | 53   | 0      | 1           | 0                             | 0              | 0              | 1                      | 168    | 76.8   | 27.2                  | 42.6                   | 18.1                      | 24.5                      | 42                                             | 22.3                            | 77.73              | 17.45                              | 11                        | 3              | 56.7             | 23.7        | 17.1                    | 9.8                     | 31.7                            | 3                              | 127     | 128                    | 78                            | 41                             | 130          | 27                             | 38                           | 50                         | 0.9        | 104.05                           | 0                        | 16.1             | 30.9
1                | 51   | 1      | 0           | 0                             | 0              | 0              | 0                      | 163    | 79.1   | 29.8                  | 34.1                   | 13                        | 21                        | 38.12                                          | 38.43                           | 61.57              | 14.64                              | 9                         | 2.5            | 46.3             | 21.51       | 30.4                    | 15.16                   | 27.3                            | 0                              | 92      | 183                    | 96                            | 70                             | 63           | 16                             | 18                           | 109                        | 0.64       | 106.93                           | 1.7                      | 14               | 8.2
1                | 53   | 1      | 0           | 0                             | 0              | 0              | 0                      | 164    | 78.5   | 29.2                  | 36.2                   | 16.4                      | 19.8                      | 45                                             | 35.2                            | 64.84              | 14.71                              | 8                         | 2.6            | 48.3             | 32.6        | 27.6                    | 12.6                    | 28.2                            | 2                              | 99      | 183                    | 93                            | 73                             | 103          | 14                             | 19                           | 60                         | 0.6        | 107                              | 2.4                      | 12.8             | 17
1                | 59   | 1      | 0           | 0                             | 0              | 0              | 0                      | 155    | 70.7   | 29.4                  | 33.9                   | 13                        | 21                        | 38.34                                          | 31.12                           | 68.88              | 16.66                              | 7                         | 2.5            | 46.2             | 15.39       | 22                      | 9.14                    | 27.3                            | 0                              | 96      | 245                    | 174                           | 53                             | 124          | 11                             | 15                           | 108                        | 0.83       | 80                               | 3.7                      | 15.6             | 6.6
0                | 35   | 1      | 0           | 0                             | 0              | 0              | 0                      | 165    | 66.3   | 24.4                  | 34.8                   | 14.8                      | 20                        | 43                                             | 26.7                            | 73.3               | 15.46                              | 3                         | 2.5            | 46.1             | 10.7        | 17.7                    | 7.9                     | 26.6                            | 0                              | 94      | 122                    | 57                            | 55                             | 34           | 15                             | 16                           | 64                         | 0.7        | 115.59                           | 0                        | 11.9             | 26.9
0                | 56   | 0      | 0           | 0                             | 0              | 0              | 0                      | 184    | 105.9  | 31.3                  | 53.1                   | 22.2                      | 30.9                      | 42                                             | 28                              | 71.95              | 17.35                              | 15                        | 3.7            | 72.5             | 42.1        | 29.7                    | 18.2                    | 40.4                            | 2                              | 114     | 267                    | 167                           | 29                             | 295          | 20                             | 35                           | 92                         | 1.07       | 81.44                            | 0.08                     | 15.7             | 27.3
... (47 rows omitted)
```
failure_message: "❌ validation does not seem correct based on this random seed value."
success_message: "✅ validation seems correct based on this random seed value."



Tests for task_07 (points: None)
```
>>> isinstance(train_transformed, Table)
True
```
failure_message: "❌ train_transformed is not assigned to a Table."
success_message: "✅ train_transformed is assigned to a Table."

```
>>> isinstance(validation_transformed, Table)
True
```
failure_message: "❌ validation_transformed is not assigned to a Table."
success_message: "✅ validation_transformed is assigned to a Table."

```
>>> isinstance(test_transformed, Table)
True
```
failure_message: "❌ test_transformed is not assigned to a Table."
success_message: "✅ test_transformed is assigned to a Table."

```
>>> train_transformed
Gallstone Status | Vitamin D_standardized | C-Reactive Protein (CRP)_standardized | Total Body Water (TBW)_standardized | Lean Mass (LM) (%)_standardized
0                | -1.47226               | -0.692676                             | -0.562662                           | -0.821138
0                | 0.258629               | -0.692676                             | 0.392375                            | 0.0501919
1                | -1.5466                | 0.322572                              | -0.513049                           | -0.415786
1                | -0.909458              | 2.72225                               | 1.3102                              | -0.509694
0                | 1.09222                | 0.885573                              | -0.612274                           | -0.37418
0                | 1.42672                | 0.507163                              | 0.069895                            | 0.503093
0                | 0.304139               | -0.185052                             | 0.652839                            | -0.370614
1                | -0.229844              | -0.692676                             | 0.144313                            | 1.19612
1                | 1.46919                | 0.461015                              | -0.19057                            | 1.12598
0                | 0.816125               | -0.397331                             | 0.752064                            | 0.721817
... (160 rows omitted)
```
failure_message: "❌ train_transformed does not seem correct based on this random seed"
success_message: "✅ train_transformed seems correct based on this random seed."

```
>>> validation_transformed
Gallstone Status | Vitamin D_standardized | C-Reactive Protein (CRP)_standardized | Total Body Water (TBW)_standardized | Lean Mass (LM) (%)_standardized
0                | 1.00833                | -0.692676                             | 1.02493                             | 1.95927
1                | 0.258629               | -0.692676                             | -1.48049                            | 0.586304
0                | 0.683388               | -0.646528                             | -0.227779                           | 0.501905
1                | -1.00503               | -0.692676                             | -0.947157                           | -1.19083
0                | 0.991338               | -0.692676                             | 0.16912                             | 0.744403
1                | -1.41917               | 0.876344                              | -0.885142                           | -1.17657
1                | -0.484699              | 1.52241                               | -0.624677                           | -0.787854
1                | -1.58907               | 2.72225                               | -0.909948                           | -0.307612
0                | 0.566579               | -0.692676                             | -0.79832                            | 0.217801
0                | 0.609055               | -0.61884                              | 1.47144                             | 0.0573242
... (47 rows omitted)
```
failure_message: "❌ validation_transformed does not seem correct based on this random seed"
success_message: "✅ validation_transformed seems correct based on this random seed."

```
>>> test_transformed
Gallstone Status | Vitamin D_standardized | C-Reactive Protein (CRP)_standardized | Total Body Water (TBW)_standardized | Lean Mass (LM) (%)_standardized
1                | -1.18555               | 5.30652                               | -1.04638                            | -2.09188
0                | 0.144475               | -0.2312                               | -0.178166                           | 1.58007
0                | 0.322343               | -0.194282                             | -0.885142                           | -1.31565
0                | -1.46164               | -0.692676                             | -0.909948                           | -1.24789
0                | 0.24801                | -0.692676                             | 1.17377                             | 0.146478
0                | 1.46919                | -0.692676                             | 0.0202827                           | 1.51231
0                | 1.10815                | -0.692676                             | -1.48049                            | 0.377089
1                | -1.36607               | 1.06093                               | -1.76576                            | -0.0235086
0                | 1.42672                | -0.692676                             | 0.069895                            | 1.49686
0                | 0.364819               | -0.692676                             | 0.231135                            | -0.507317
... (54 rows omitted)
```
failure_message: "❌ test_transformed does not seem correct based on this random seed"
success_message: "✅ test_transformed seems correct based on this random seed."



Tests for task_08 (points: None)
```
>>> callable(classify)
True
```
failure_message: "❌ classify is not callable."
success_message: "✅ classify is callable."

```
>>> training_ex = Table().with_columns('Gallstone Status', [1, 0, 1, 1, 0], 'Feature 1', [3, 13, 4, 3.4, 11], 'Feature 2', [75, 78, 71, 70, 79])
>>> classify(training_ex, [12, 76], 3) in {0, 1}
True
```
failure_message: "❌ classify didn't produce a value of 0 or 1 for an example data set.
success_message: "✅ classify produced a value of 0 or 1 for an example data set."

```
>>> training_ex = Table().with_columns('Gallstone Status', [1, 0, 1, 1, 0], 'Feature 1', [3, 13, 4, 3.4, 11], 'Feature 2', [75, 78, 71, 70, 79])
>>> classify(training_ex, [12, 76], 3)
0
```
failure_message: "❌ classify doesn't seem to be working correctly based on an example data set."
success_message: "✅ classify seems to be working correctly based on an example data set."

```
>>> an_example = [-0.03995, -0.68226, -0.9853, 1.10132]
>>> classify(train_transformed, an_example, 5)
0
```
failure_message: "❌ classify doesn't seem to be working correctly based on an example from the validation set"
success_message: "✅ classify seems to be working correctly based on an example from the validation set."



Tests for task_09 (points: None)
```
>>> callable(evaluate_accuracy)
True
```
failure_message: "❌ evaluate_accuracyy is not callable."
success_message: "✅ evaluate_accuracy is callable."

```
>>> evaluate_accuracy(train_transformed, validation_transformed, 3)
0.631578947368421
```
failure_message: "❌ evaluate_accuracy doesn't seem to be working for the train and validation sets."
success_message: "✅ evaluate_accuracy seems to be working for the train and validation sets."



Tests for task_11 (points: None)
```
>>> import numbers
>>> isinstance(classifier_accuracy, numbers.Number)
True
```
failure_message: "❌ classifier_accuracy is not assigned to a numerical value."
success_message: "✅ classifier_accuracy is assigned to a numerical value."

```
>>> import hashlib
>>> 
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(round(classifier_accuracy, 3))
'5d8da4f6668f40f340a752537beb1a4f'
```
failure_message: "❌ classifier_accuracy does not seem correct for this random seed."
success_message: "✅ classifier_accuracy seems correct for this random seed."

