<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Guided Learning Activity 10: Central Limit Theorem

This Guided Learning Activity is designed for you to complete alongside someone knowledgable of MATH 108 Foundations of Data Science at CCSF or an equivalent course at another institution. Keep in mind that this is not a graded assignment for MATH 108 by default. If you have any concerns about participation, reach out to your instructor.

---

## Learning Objectives

1. Summarize the centrality and variability of a data set using a variety of measurements.
2. Explain how sample size is related to the variability of the sample mean.
3. Use the Central Limit Theorem to perform generate a confidence interval esiamte for a relavent parameter.

---

## Configure the Notebook

### Task 00 📍

<!-- BEGIN QUESTION -->

Run the following code cell to set up this notebook for this activity.

```python
import os
import otter
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
from scipy import stats
from statistics import mode
import pandas as pd
!pip install -q xlrd

try:
    notebook_name = "gla10.ipynb"
    session_name = os.getenv('JPY_SESSION_NAME')
    
    if session_name:
        notebook_name = os.path.basename(session_name)
        print(f"Detected notebook name from environment: {notebook_name}")
    else:
        print("Environment variable 'JPY_SESSION_NAME' not found. Using default notebook name: gla01.ipynb")
        
    grader = otter.Notebook(notebook_name)
    print(f"The Otter grader has been initialized with {notebook_name}.")
except Exception as e:
    raise RuntimeError(f"Failed to initialize Otter grader: {e}") from e
```

<!-- END QUESTION -->

---

---

## Overview

---

### Measures of Centrality and Variability

* It is impossible to summarize data with a single value.
* Every summary involves some loss of information.
* It is common to describe the distribution from two perspectives:
    * Centrality Measures – describe the center of the data.
        * Examples: Mean, Median, Mode
    * Variability Measures – describe how the data is spread.
        * Examples: Range, Interquartile Range, Variance, Standard Deviation, Number of Categories

---

#### Examples Metrics

| **Type** | **Measure**                   | **Description**                                                       | **Recommended Use**                                                            |
| ------------ | ----------------------------- | --------------------------------------------------------------------- | -------------------------------------------------------------------------- |
| **Center**   | **Mean or Average**            | Sum of all values divided by the number of values.                    | When data are roughly symmetric and have no extreme outliers.              |
| **Center**   | **Median**                    | Middle value when data are ordered.                                   | When data are skewed or have outliers.                                     |
| **Center**   | **Mode**                      | Most frequently occurring value.                                      | When data are categorical or have repeated values.                         |
| **Spread**   | **Range**                     | Difference between maximum and minimum values.                        | Quick sense of total spread; sensitive to outliers.                        |
| **Spread**   | **Interquartile Range (IQR)** | Difference between the 75th and 25th percentiles (Q3 − Q1).           | When comparing variability of skewed data or data with outliers.           |
| **Spread**   | **Variance**                  | Average squared deviation from the mean.                              | When analyzing data mathematically (used in formulas).                     |
| **Spread**   | **Standard Deviation (SD)**   | Square root of the variance; measures average distance from the mean. | When data are roughly symmetric and you want spread in same units as data. |
| **Spread**   | **Number of Categories**   | The number of categories. | When data are categorical. |


---

#### Properties of Means

- From the [MATH 108 textbook](https://ccsf-math-108.github.io/textbook/chapters/14/1/Properties_of_the_Mean.html):
    - It need not be an element of the collection.
    - It need not be an integer even if all the elements of the collection are integers.
    - It is somewhere between the smallest and largest values in the collection.
    - It need not be halfway between the two extremes; it is not in general true that half the elements in a collection are above the mean.
    - If the collection consists of values of a variable measured in specified units, then the mean has the same units too.
- The mean value is the value which equializes, balances the entire distribution. 
- A proportion is a mean.
- In relation to the medain:
    - The mean and median are equal in a perfectly symetric distribution.
    - The mean is larger (smaller) than the median in a distribuiton that is skewed right (left).

---

### Chebychev’s Bound

- For **all** collections of numerical values and all numbers $z$, the proportion of entries that are in the range “average $\pm z$ SDs” is at least $1 - \frac{1}{z^2}$.
- Any numerical data value can be described in standard unit as $z = \frac{\text{value} - \text{average}}{\text{SD}}$.

---

### The Normal Distribution

---

#### Overview

- The normal curve appears a lot in data/statistical analysis.
- Each normal curve is uniquely defined by its mean $\mu$ and standard deviation $\sigma$.
- The standard normal curve has a mean value of 0 and a standard deviation of 1.
- There are many situations where a histogram's outline can be approximated with a normal curve.
- Like a histogram height, the densities (heights) associated with a normal curve is calculated by `stats.norm.pdf` from the `scipy` library.
- Like the areas of the bars in a histogram, the area under the the normal curve is calculated by `stats.norm.cdf` from the `scipy` library.
- The inflection points on the normal curve help identify 1 standard deviation away from the mean (center).

---

#### Visualizing Areas

- When discussing normally distributed data (or data that is appoximately normal), there are useful esimates that can be made with the mean and standard deviation.
- Three common estimates are visaulized below:

<img src="./normal_approx.png" alt="3 Normal curves showing the data approximations related 1, 2, and 3 standard deviations within the mean.">

---

### The Central Limit Theorem

- How does the Normal distribution relate to summary statistics like averages and other sums?
- Regardless of the distribution of a variable, the probability distribution of the sum or average of a large random sample drawn with replacement will be roughly normal.
- For any population, the distribution of very possible sample average can be described b:
    - The mean of the probability distribution for a sample proportion is the mean of the population.
    - The standard deviation of the probability distribution for a sample average is the population standard deviation divided by the square root of the sample size.
- This means that you can create a confidence interval approxiamation for a mean using $$\mu \pm z \frac{\sigma}{\sqrt{n}}$$                                                                 
    - $n$ is the sample size
    - $\mu$ is usually approximated by the average of the same
    - $\sigma$ is estimated too, but that estiamation varies.
    - the confidence level is associated with the value of $z$ (e.g. $z=2$ for 95\% confidence)

---

### Sample Size

- A 95\% confidence interval estimate for the population mean can be expressed as $\mu \pm 2 \frac{\sigma}{\sqrt{n}}$.
- The width of the interval is $4 \frac{\sigma}{\sqrt{n}}$.
- If you want the width to be at most a certain value $W$ and the population is made of binary values, then the sample size needs to be at least $\left(4 \cdot \frac{0.5}{W}\right) ^2$.

---

## The Superstore

---

### The Data

<a href="https://unsplash.com/photos/a-couple-of-men-standing-next-to-each-other-in-a-room-SB_LIDxySPY"><img src="premium_photo-1739350938199-89ac8ecacb26.png" width=600px alt="Two people shopping in a store for some furniture"/></a>

The data in the `XLS` file `Sample - Superstore.xls` represents a dataset from [Tableau](https://help.tableau.com/current/guides/get-started-tutorial/en-us/get-started-tutorial-connect.htm) for a ficticious company called Superstore that provides data about sales, products, customers, and geography. Run the following cell to assign the data to the table `superstore` using the Pandas function `pd.read_excel`.

```python
df = pd.read_excel('Sample - Superstore.xls')
superstore = Table.from_df(df).drop(0)
superstore.show(3)
```

### Task 01 📍🔎

<!-- BEGIN QUESTION -->

Visualize the distribution of sub-categories in `superstore`, and provide an appropriate measure of central and variability.

_Type your answer here, replacing this text._

```python
# Create your visualization
...
```

<!-- END QUESTION -->

---

### Task 02 📍🔎

<!-- BEGIN QUESTION -->

Visualize the distribution of profits and provide an appropriate measure of central and variability.

_Type your answer here, replacing this text._

```python
# Create your visualization
...
```

<!-- END QUESTION -->

---

### Task 03 📍🔎

<!-- BEGIN QUESTION -->

If you ignore the bottom 5% and top 5% of the profit data in `superstore`, would you change you summarize the distribution any differently?

_Type your answer here, replacing this text._

```python
# Create your visualization
...
```

<!-- END QUESTION -->

---

### Inference Goal

Let’s say you’re an analyst working for Superstore, and your goal is to figure out what proportion of corporate orders in California will yield a profit greater than \$20. 

* You could calculate the proportion from the data you already have, but that number probably wouldn’t be exactly right for all the corporate orders this made-up company might get in the future.
* A better approach might be to build a confidence interval using the bootstrap resampling method you learned earlier. That would give you a range that will likely capture the true proportion.
* Alternately, you could also use the Central Limit Theorem to help with this. It’s a useful tool for estimating sample sizes and building confidence intervals when you're dealing with sums, which a proportion is a type of sum.

---

### Task 04 📍

Create a table called `cal_corp` which contains the data in `superstore` for California corporate orders, and add a column called `'Above $20 Profit` which contains binary values `bool` indicating whether the order yielded a profit about \$20.

```python
...
```

```python
grader.check("task_04")
```

---

### Task 05 📍

What proportion of the California corporate orders yeilded a profit above \$20? Assign that proportion to the variable `above_20_prop`. The proportion should be a numerical value between 0 and 1, inclusive.

```python
above_20_prop = ...
above_20_prop
```

```python
grader.check("task_05")
```

---

### Task 06 📍

What is the smallest sample size needed to ensure that an interval esimate for the proportion of California corporate orders with a profit above \$20 is within 0.1 for a confidence level of 95\%? Assign that sample size (`int`)to the variable `sample_size`.

```python
...
```

```python
grader.check("task_06")
```

---

### Task 07 📍🔎

<!-- BEGIN QUESTION -->

Create a 95\% confidence interval estimate for the proportion using the Central Limit Theorem with the sample proportion as an estimate for $\mu$, $n=618$, and $0.5$ for an esimate of $\mu$. Assign the lower bound of the interval to `lower_bound` and the upper bound of the interval to `upper_bound`.

```python
lower_bound = ...
upper_bound = ...
print(f'The 95% confidence interval for the proportion is ({lower_bound:.4f}, {upper_bound:.4f})')
```

```python
grader.check("task_07")
```

<!-- END QUESTION -->

---

### Task 08 📍

Create Create a 95\% confidence interval estimate for the proportion using bootstrap resamples. Assign the lower bound of the interval to `lower` and the upper bound of the interval to `upper`.

```python
np.random.seed(2025) # Keep this so the auto-grader marks your answer correctly.

def simulate_and_calculate_prop():
    bootstrap_sample = cal_corp.sample()
    return np.mean(bootstrap_sample.column('Above $20 Profit'))

simulated_props = make_array()

for __ in np.arange(1_000):
   simulated_props = np.append(simulated_props, simulate_and_calculate_prop())
    
Table().with_column('Simulated Proportion', simulated_props).hist()

lower = ...
upper = ...
print(f'The approximate 95% confidence interval for the proportion is ({lower:.4f}, {upper:.4f})')
```

```python
grader.check("task_08")
```

---

### Task 09 📍🔎

<!-- BEGIN QUESTION -->

How do the confidence intervals in Task 07 and Task 08 compare? Also, connect these intervals to the analyts' goals. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Reflection

In this activity, you reviewed about how to describe the centrality and variability of a distribution. In particular, you focused on why the the mean and standard deviation are so important from a theoretical pespective. You saw how that theory could be used to create confidence intervals and calulate sample sizes to ensure certain error bounds. Lastly, you compared the results of generating a confidence interval from a normal distribuiton-based formula and from the bootstrap resampling technique.

---

## License

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a>.

<img src="./by-nc-sa.png" width=100px>

Tests for task_04 (points: None)
```
>>> isinstance(cal_corp, Table)
True
```
failure_message: "❌ cal_corp is not assigned to a Table."
success_message: "✅ cal_corp is assigned to a Table."

```
>>> cal_corp.num_rows
618
```
failure_message: "❌ cal_corp does not have the correct number of rows."
success_message: "✅ cal_corp has the correct number of rows."

```
>>> cal_corp.labels == ('Order ID', 'Order Date', 'Ship Date', 'Ship Mode', 'Customer ID', 'Customer Name', 'Segment', 'Country/Region', 'City', 'State', 'Postal Code', 'Region', 'Product ID', 'Category', 'Sub-Category', 'Product Name', 'Sales', 'Quantity', 'Discount', 'Profit', 'Above $20 Profit')
True
```
failure_message: "❌ cal_corp does not have the correct columns."
success_message: "✅ cal_corp has the correct column labels."

```
>>> cal_corp
Order ID       | Order Date          | Ship Date           | Ship Mode      | Customer ID | Customer Name    | Segment   | Country/Region | City          | State      | Postal Code | Region | Product ID      | Category        | Sub-Category | Product Name                                              | Sales  | Quantity | Discount | Profit  | Above $20 Profit
CA-2020-138688 | 2020-06-12 00:00:00 | 2020-06-16 00:00:00 | Second Class   | DV-13045    | Darrin Van Huff  | Corporate | United States  | Los Angeles   | California | 90036       | West   | OFF-LA-10000240 | Office Supplies | Labels       | Self-Adhesive Address Labels for Typewriters by Universal | 14.62  | 2        | 0        | 6.8714  | False
CA-2020-101343 | 2020-07-17 00:00:00 | 2020-07-22 00:00:00 | Standard Class | RA-19885    | Ruben Ausman     | Corporate | United States  | Los Angeles   | California | 90049       | West   | OFF-ST-10003479 | Office Supplies | Storage      | Eldon Base for stackable storage shelf, platinum          | 77.88  | 2        | 0        | 3.894   | False
CA-2020-109806 | 2020-09-17 00:00:00 | 2020-09-22 00:00:00 | Standard Class | JS-15685    | Jim Sink         | Corporate | United States  | Los Angeles   | California | 90036       | West   | OFF-AR-10004930 | Office Supplies | Art          | Turquoise Lead Holder with Pocket Clip                    | 20.1   | 3        | 0        | 6.633   | False
CA-2020-109806 | 2020-09-17 00:00:00 | 2020-09-22 00:00:00 | Standard Class | JS-15685    | Jim Sink         | Corporate | United States  | Los Angeles   | California | 90036       | West   | TEC-PH-10004093 | Technology      | Phones       | Panasonic Kx-TS550                                        | 73.584 | 2        | 0.2      | 8.2782  | False
CA-2020-109806 | 2020-09-17 00:00:00 | 2020-09-22 00:00:00 | Standard Class | JS-15685    | Jim Sink         | Corporate | United States  | Los Angeles   | California | 90036       | West   | OFF-PA-10000304 | Office Supplies | Paper        | Xerox 1995                                                | 6.48   | 1        | 0        | 3.1104  | False
CA-2021-106180 | 2021-09-18 00:00:00 | 2021-09-23 00:00:00 | Standard Class | SH-19975    | Sally Hughsby    | Corporate | United States  | San Francisco | California | 94122       | West   | OFF-AR-10000940 | Office Supplies | Art          | Newell 343                                                | 8.82   | 3        | 0        | 2.3814  | False
CA-2021-106180 | 2021-09-18 00:00:00 | 2021-09-23 00:00:00 | Standard Class | SH-19975    | Sally Hughsby    | Corporate | United States  | San Francisco | California | 94122       | West   | OFF-EN-10004030 | Office Supplies | Envelopes    | Convenience Packs of Business Envelopes                   | 10.86  | 3        | 0        | 5.1042  | False
CA-2021-106180 | 2021-09-18 00:00:00 | 2021-09-23 00:00:00 | Standard Class | SH-19975    | Sally Hughsby    | Corporate | United States  | San Francisco | California | 94122       | West   | OFF-PA-10004327 | Office Supplies | Paper        | Xerox 1911                                                | 143.7  | 3        | 0        | 68.976  | True
CA-2019-124919 | 2019-05-31 00:00:00 | 2019-06-02 00:00:00 | First Class    | SP-20650    | Stephanie Phelps | Corporate | United States  | San Jose      | California | 95123       | West   | OFF-PA-10001950 | Office Supplies | Paper        | Southworth 25% Cotton Antique Laid Paper & Envelopes      | 58.38  | 7        | 0        | 26.271  | True
CA-2019-124919 | 2019-05-31 00:00:00 | 2019-06-02 00:00:00 | First Class    | SP-20650    | Stephanie Phelps | Corporate | United States  | San Jose      | California | 95123       | West   | OFF-PA-10002254 | Office Supplies | Paper        | Xerox 1883                                                | 105.52 | 4        | 0        | 48.5392 | True
... (608 rows omitted)
```
failure_message: "❌ cal_corp does not seem correct."
success_message: "✅ cal_corp seems correct."



Tests for task_05 (points: None)
```
>>> from numbers import Number
>>> isinstance(above_20_prop, Number) and 0 <= above_20_prop <= 1
True
```
failure_message: "❌ yfinance is not a numerical value between 0 and 1, inclusive."
success_message: "✅ above_20_prop is a numerical value between 0 and 1, inclusive."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(round(above_20_prop, 3))
'b2adb6eea82ff2faa29d67e260ad703a'
```
failure_message: "❌ above_20_prop is not assigned to a reasonably correct proportion."
success_message: "✅ above_20_prop is assigned to a reasonably correct proportion."



Tests for task_06 (points: None)
```
>>> isinstance(sample_size, (int, np.integer))
True
```
failure_message: "❌ sample_size is not assigned to an integer."
success_message: "✅ sample_size is assigned to an integer."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(sample_size)
'18d8042386b79e2c279fd162df0205c8'
```
failure_message: "❌ ssample_size is not assigned to the smallest sample sized needed for this situation."
success_message: "✅ sample_size is assigned to the smallest sample size needed for this situation."



Tests for task_07 (points: None)
```
>>> from numbers import Number
>>> isinstance(lower_bound, Number) and isinstance(upper_bound, Number)
True
```
failure_message: "❌ lower_bound and upper_bound are not assigned to numerical values."
success_message: "✅ lower_bound and upper_bound are assigned to numerical values."

```
>>> 0 <= lower_bound < upper_bound <= 1
True
```
failure_message: "❌ lower_bound is not less than upper_bound or they both are not between 0 and 1, inclusive."
success_message: "✅ lower_bound is less than upper_bound and they both are between 0 and 1, inclusive."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(round(lower_bound, 2))
'149dd5056939405870c9bb50cbc8691c'
```
failure_message: "❌ lower_bound does not seem reasonably correct."
success_message: "✅ lower_bound seems reasonably correct."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(round(upper_bound, 2))
'91afec64e32d6bf957e441df2ab638bb'
```
failure_message: "❌ upper_bound does not seem reasonably correct."
success_message: "✅ upper_bound seems reasonably correct."



Tests for task_08 (points: None)
```
>>> from numbers import Number
>>> isinstance(lower, Number) and isinstance(upper, Number)
True
```
failure_message: "❌ lower and upperare not assigned to numerical values."
success_message: "✅ lower and upper are assigned to numerical values."

```
>>> 0 <= lower < upper <= 1
True
```
failure_message: "❌ lower is not less than upper or they both are not between 0 and 1, inclusive."
success_message: "✅ lower is less than upper and they both are between 0 and 1, inclusive."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(round(lower, 2))
'8ce3fac7e23a02ab4e00cf0f1e03310a'
```
failure_message: "❌ lower does not seem reasonably correct."
success_message: "✅ lower seems reasonably correct."

```
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> get_hash(round(upper, 2))
'964e2b882801bd4ba988904454316d76'
```
failure_message: "❌ upper does not seem reasonably correct."
success_message: "✅ upper seems reasonably correct."

