```python
# Initialize Otter
import otter
grader = otter.Notebook("gla12.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Guided Learning Activity 12: Residual Analysis and Prediction Intervals

This Guided Learning Activity is designed for you to complete alongside a Data Ambassador from the course. You might find that it feels like a combination of the lectures and lab assignment. Whether you are participating live or watching the recording of the live meeting, let the Data Ambassador guide you through the following tasks. There will be moments for you to reflect and explore your own ideas as a way to solidify concepts and skills introduced by your instructor. Keep in mind that this is not a graded assignment for MATH 108 by default. If you have any concerns about participation, reach out to your instructor.

---

## Learning Objectives

1. Review linear regression and residuals
2. Describe how to decide between two models based on data trends, MSE values, and residual plots
3. Fit a model to real data after considering residual plots
4. Create a prediction interval estimate based on a model

---

## Configure the Notebook

Run the following code cell to set up the notebook.

```python
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')

# Functions from MATH 108 textbook
def standard_units(any_numbers):
    """Convert any array of numbers to standard units."""
    return (any_numbers - np.mean(any_numbers))/np.std(any_numbers)  

def correlation(t, label_x, label_y):
    return np.mean(standard_units(t.column(label_x))*standard_units(t.column(label_y)))

def slope(t, label_x, label_y):
    r = correlation(t, label_x, label_y)
    return r*np.std(t.column(label_y))/np.std(t.column(label_x))

def intercept(t, label_x, label_y):
    return np.mean(t.column(label_y)) - slope(t, label_x, label_y)*np.mean(t.column(label_x))

def fit(table, x, y):
    """Return the height of the regression line at each x value."""
    a = slope(table, x, y)
    b = intercept(table, x, y)
    return a * table.column(x) + b

def residual(table, x, y):
    return table.column(y) - fit(table, x, y)
```

---

## Residual Analysis

---

### Linear Regression Review

- The **least squares regression line** minimizes the **mean squared error (MSE)**.
- It's the best-fitting **line**, but only if a line makes sense.

---

### Residuals

<img src="./residuals_annotated.png" width=700px alt='Residuals annotated for a few data points with the regression line shown'>

- A residual (error) is the difference between the actual value and the predicted value.
- The average residual is 0
- The SD of the residuals is a fraction of the SD of $y$
    - The fraction is $\sqrt{1 - r^2}$
        - $r$ is the correlation coefficient for $x$ and $y$
    - An interpretation of $r^2$: The proportion of variance in the $y$ that is predicted by the $x$.
        - $r^2$ (or $R^2$ in common notation) is known as the Coefficient of Determination
        - Higher $R^2$ indicates a better fit of the model to the data

---

### When is a Line a Good Choice?

- Visual Inspection:
    - Create a **scatterplot** of the variables
    - If the trend is **linear**, a line may work well
    - If the trend is **curved or uneven**, consider a **nonlinear** model
- Compare Models:
    - Fit several reasonable models (e.g., linear, polynomial, log)
    - Compare their **MSE** values — lower is better
- Examine Residual Plots:
    - A residual plot is a scatter plot of $y$ vs. residuals
    - If residuals are **randomly scattered** (no pattern), a linear model is likely appropriate
    - If residuals show a curve, funnel, or pattern, then consider another model

---

### Task 01 📍🔎

<!-- BEGIN QUESTION -->

The following graphic shows the scatterplot with regression line and MSE along with the associated residual plot for two relationships. Which relationship is best suited for a linear regression model? Address the way the regression line fits the data in the scatterplot, compare the MSE values, and compare the residuals plots to support your opinion.

<img src="./two_relationships.png" width=800px alt="Scatterplot with regression line and MSE and residual plot for two relationships">

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Michaelis–Menten Kinetics

<a href="https://en.wikipedia.org/wiki/Michaelis%E2%80%93Menten_kinetics" target="_blank"><img src="./MM-curve.jpg" width=400px alt="Curve of the Michaelis–Menten equation labelled in accordance with IUBMB recommendations"></a>

From <a href="https://en.wikipedia.org/wiki/Michaelis%E2%80%93Menten_kinetics" target="_blank">Wikipedia</a>:

> In biochemistry, Michaelis–Menten kinetics, named after Leonor Michaelis and Maud Menten, is the simplest case of enzyme kinetics, applied to enzyme-catalysed reactions involving the transformation of one substrate into one product.

---

### Biochemistry Basics

- **Enzyme**: A protein that speeds up chemical reactions in the body.
- **Substrate**: The molecule that the enzyme acts on (like a key in a lock).
- **Enzyme + substrate → product**: The reaction rate depends on substrate amount.
- **Saturation point**: When all enzymes are busy, adding more substrate won't increase speed.

---

### The Data

The file [`puromycin.csv`](https://www.rdocumentation.org/packages/datasets/versions/3.6.2/topics/Puromycin) contains 23 rows and 3 columns of the reaction velocity versus substrate concentration in an enzymatic reaction involving untreated cells or cells treated with [Puromycin](https://en.wikipedia.org/wiki/Puromycin), an inhibitor that slows down enzymatic reactions.

---

### Task 02 📍

Assign the `puromycin.csv` to the `Table` called `puromycin`. The table should contain only the untreated data and the following columns:

- `'Concentration'`: a numeric vector of substrate concentrations (ppm)
- `'Reaction Rate'`: a numeric vector of instantaneous reaction rates (counts/min/min)

```python
puromycin = ...
puromycin
```

```python
grader.check("task_02")
```

---

### Task 03 📍🔎

<!-- BEGIN QUESTION -->

Visualize the relations between the substrate concentrations and reaction rates, where the concentrations are the explanatory variable. Include the best fit line in the visualization.

```python
...
plt.title('Reaction Rate vs. Substrate Concentration')
plt.show()
```

<!-- END QUESTION -->

---

### Task 04 📍

Create a table called `puromycin_with_residuals` that contains the same information from `puromycin` with an additional column called `Residual` that contains the residual based on a linear regression model.

**Hint**: Utilize the `residual` function from MATH 108.

```python
puromycin_with_residuals = ...
puromycin_with_residuals
```

```python
grader.check("task_04")
```

---

### Task 05 📍🔎

<!-- BEGIN QUESTION -->

Create a residual plot for the linear model associated with the concentration and reaction rate data. Based on the residual plot, should we use a linear model for this data?

_Type your answer here, replacing this text._

```python
...
plt.title('Residual Plot')
plt.show()
```

<!-- END QUESTION -->

---

### A Better Model

The key formula for this relationship is called the Michaelis-Menten equation: 
  $$
  v = \frac{V_{max} \cdot [S]}{K_m + [S]}
  $$ 
- $v$: reaction rate
- $V_{max}$: maximum reaction rate achieved
- $[S]$: substrate concentration
- $K_m$: Michaelis constant (the concentration at half the maximum reaction rate)

---

### Task 06 📍

Using the data in `puromycin`, fit a function based on the Michaelis–Menten equation by minimizing MSE. Since the formula depends on the values of the parameters $V_{max}$ and $K_m$, this means to use the `minimize` function to find approximate values for those parameters. Assign those approximate values to `V_max` and `K_m`.

**Hints**: 
- Follow the [nonlinear regression example](https://ccsf-math-108.github.io/textbook/chapters/15/4/Least_Squares_Regression.html#nonlinear-regression) from the MATH 108 textbook.
- We've provided you with a template for the `MM` function (Michaelis–Menten equation) and the MSE function (`MM_MSE`) for an Michaelis–Menten equation-based model.

```python
def MM(S, V_max, K_m):
    return V_max * S / (K_m + S)

def MM_MSE(V_max, K_m):
    S = puromycin.column('Concentration')
    v = puromycin.column('Reaction Rate')
    fitted_v = MM(S, V_max, K_m)
    return np.mean((v - fitted_v) ** 2)

V_max, K_m = minimize(MM_MSE)
print(f'Maximum Reaction Rate: {V_max}\n\
Michaelis constant: {K_m}')
```

```python
def MM(S, V_max, K_m):
    return ...

def MM_MSE(V_max, K_m):
    S = ...
    v = ...
    fitted_v = ....
    return ...

V_max, K_m = minimize(MM_MSE)
print(f'Maximum Reaction Rate: {V_max}\n\
Michaelis constant: {K_m}')
```

```python
grader.check("task_06")
```

---

### Task 07 📍🔎

<!-- BEGIN QUESTION -->

Now that you've fitted the new model to the data, create your own `fit` and `residual` functions `fit_MM` and `residual_MM` to incorporate the `MM` function. We've provided you with a template.

**Notes**: 
- You'll see a plot of the data with the associated predictions based on the equation.
- Although the residual plot may appear heteroscedastic ('fanning' out), this effect may be due to limited data at high concentration levels, not necessarily a poor model fit.

```python
def fit_MM(table, S_label, V_max, K_m):
    '''Return the predicted reaction rate based on
    the Michaelis-Menten equation at each concentration value.'''
    S = table.column(S_label)
    return ...

def residual_MM(table, S_label, v_label, V_max, K_m):
    '''Calculate the residulas based on fit_MM'''
    return ....

# Create a table for plotting
puromycin_for_plotting = puromycin.with_columns(
    'Predicted Reaction Rate', 
    fit_MM(puromycin, 'Concentration', V_max, K_m),
    'Residual', 
    residual_MM(puromycin, 'Concentration', 
                'Reaction Rate', V_max, K_m))

# Graph Predictions
puromycin_for_plotting.scatter('Concentration', 
                               ['Reaction Rate', 
                                'Predicted Reaction Rate'])
plt.title('Concentration vs. Reaction Rate with Predictions')
plt.show()

# Residual Plot
puromycin_for_plotting.scatter('Concentration', 'Residual')
plt.title('Residual Plot')
plt.show()
```

<!-- END QUESTION -->

---

### Prediction Intervals

- When you create a prediction model, you are likely using sample data to do so
- If you had different sample data, then you'll likely create a different model
- The predictions for the same value from different models will be different
- So, a single prediction is very likely to be wrong, but an interval of predictions is likely to capture the true value
- Resampling the available data would provide you with a way to create new predictions
- The following graph shows an original set of data with a linear model along with a resampled set of that original data with the updated linear model.
- Notice the slight difference in the model lines.

<img src="./original_resampled_with_models.png" width=600px alt="Two sets of data (an original set and resampled set with linear models for each.">

---

### Task 08 📍

The following code shows that the predicted reaction rate for a concentration of 0.8 ppm is approximately 151.26 counts/min/min.

```python
predicted_v = MM(0.8, V_max, K_m)
predicted_v
```

Using our bootstrap methodology, create a 95% confidence interval estimate for the predicted reaction rate for a concentration of 0.8 ppm.

```python
...
```

```python
grader.check("task_08")
```

---

## Reflection

In this activity, you considered how residuals can help you determine whether to use a linear model or some other type of model. Additionally, you fit a nonlinear model to a real data set and created a prediction interval based on a bootstrap resampling method.

---

## License

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a>.

<img src="./by-nc-sa.png" width=100px>

Tests for task_02 (points: None)
```
>>> isinstance(puromycin, Table)
True
```
failure_message: "❌ puromycin is not assigned to a Table."
success_message: "✅ puromycin is assigned to a Table."

```
>>> puromycin.num_rows
11
```
failure_message: "❌ puromycin does not have the correct number of rows."
success_message: "✅ puromycin has the correct number of rows."

```
>>> set(puromycin.labels) == {'Concentration', 'Reaction Rate'}
True
```
failure_message: "❌ puromycin does not have the correct labels."
success_message: "✅ puromycin has the correct labels."

```
>>> puromycin
Concentration | Reaction Rate
0.02          | 67
0.02          | 51
0.06          | 84
0.06          | 86
0.11          | 98
0.11          | 115
0.22          | 131
0.22          | 124
0.56          | 144
0.56          | 158
... (1 rows omitted)
```
failure_message: "❌ puromycin does not seem correct."
success_message: "✅ puromycin seems correct."



Tests for task_04 (points: None)
```
>>> isinstance(puromycin_with_residuals, Table)
True
```
failure_message: "❌ puromycin_with_residuals is not a Table."
success_message: "✅ puromycin_with_residuals is a Table."

```
>>> puromycin_with_residuals.labels == ('Concentration', 'Reaction Rate', 'Residual')
True
```
failure_message: "❌ puromycin_with_residuals doesn't have the correct labels."
success_message: "✅ ppuromycin_with_residuals has the correct labels."

```
>>>
```
```
>>> np.all(puromycin_with_residuals.column(0) == puromycin.column(0))
True
```
failure_message: "❌ The concentration levels in puromycin_with_residuals are not the same as in puromycin."
success_message: "✅ The concentration levels in puromycin_with_residuals are the same as in puromycin."

```
>>> np.all(puromycin_with_residuals.column(1) == puromycin.column(1))
True
```
failure_message: "❌ The reactions rates in puromycin_with_residuals are not the same as in puromycin."
success_message: "✅ The reactions rates in puromycin_with_residuals are the same as in puromycin."

```
>>> puromycin_with_residuals
Concentration | Reaction Rate | Residual
0.02          | 67            | -20.8243
0.02          | 51            | -36.8243
0.06          | 84            | -7.39785
0.06          | 86            | -5.39785
0.11          | 98            | 2.13527
0.11          | 115           | 19.1353
0.22          | 131           | 25.3081
0.22          | 124           | 18.3081
0.56          | 144           | 7.93332
0.56          | 158           | 21.9333
... (1 rows omitted)
```
failure_message: "❌ puromycin_with_residuals doesn't seem correct."
success_message: "✅ puromycin_with_residuals seems correct."



Tests for task_06 (points: None)
```
>>> np.isclose(V_max, 160.28132136084938)
True
```
failure_message: "❌ V_max doesn't seem correct."
success_message: "✅ V_max seems correct."

```
>>> np.isclose(K_m, 0.047707673969251935)
True
```
failure_message: "❌ K_m doesn't seem correct."
success_message: "✅ K_m seems correct."



Tests for task_08 (points: None)
```
>>> lower > 135
True
```
failure_message: "❌ The lower bound does not seem reasonable."
success_message: "✅ The lower bound seems reasonable."

```
>>> upper < 160
True
```
failure_message: "❌ The lower bound does not seem reasonable."
success_message: "✅ The upper bound seems reasonable."

```
>>> lower < predicted_v < upper
True
```
failure_message: "❌ predicted_v should be in the interval."
success_message: "✅ predicted_v is in the interval"

