```python
# Initialize Otter
import otter
grader = otter.Notebook("gla04.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Guided Learning Activity 04: Visualizing Trends over Time

This Guided Learning Activity is designed for you to complete alongside a Data Ambassador from the course. You might find that it feels like a combination of the lectures and lab assignment. Whether you are participating live or watching the recording of the live meeting, let the Data Ambassador guide you through the following tasks. There will be moments for you to reflect and explore your own ideas as a way to solidify concepts and skills introduced by your instructor. Keep in mind that this is not a graded assignment for MATH 108 by default. If you have any concerns about participation, reach out to your instructor.

---

## Learning Objectives

1. Convert date strings into `datetime` format for proper handling in Python.
2. Create and interpret line plots to visualize time trends.
3. Identify key economic trends using CPI data from FRED and recognize anomalies.
4. Compute and visualize relative percentage changes in time-series data.

---

## Working with Dates

Many datasets include some information on date and time that need proper handling. Formatting this data as strings is generally not the best way:

- **Sorting Matters**
    - Python sorts date stored as strings alphabetically, not chronologically
    - When sorting `'February 28', 'April 2', 'April 1'`:
        - We want: `'February 28'`, `'April 1'`, `'April 2'`
        - Python gives `'April 1'`, `'April 2'`, `'February 28'`
- **Extract Useful Information**
    - Easily get the day of the week, hour, or month from a timestamp
        - `'February 28'.month` should produce `'February'`
    - Strings do you not have a `month` attribute
- **Time Differences.**
    - Calculate durations, intervals, and trends over time
    - Subtracting `12:15 PM` from `3:30 PM` should give 3 hours and 15 minutes 
- **Better Visuals.**
    - Convert timestamps to plot meaningful time-based trends

---

### `datetime` Module

* In MATH 108, you are not really expected to know how to work with dates
* It is very common to work with data that contains date information
* Python's `datetime` module helps manipulate, compare, and format dates and times
* You can view the [official documentation](https://docs.python.org/3/library/datetime.html) for more details.

---

#### Import `datetime`

Import everything from the `datetime` module.

```python
from datetime import *
```

---

#### Create `datetime` Objects

You can create a `datetime` object using `datetime(year, month, day, hour, minute, second)`.

```python
datetime(2025, 1, 31, 12, 15, 33)
```

You can even get the current date/time.

```python
datetime.now()
```

```python
datetime.today()
```

---

#### Convert `str` to `datetime`

* It is common to work with dates that are initially formatted as strings.
* This module provides `strptime` as a way to convert the strings into `datetime` objects.

```python
date_as_string = "February 21, 2025"
datetime.strptime(date_as_string, "%B %d, %Y")
```

---

#### Convert `datetime` to `str`

* Maybe you need to work with `datetime`, but the data needs to be presented/stored as a string.
* This module provides `strftime` as a way to convert the `datetime` objects into strings.

```python
datetime_object = datetime(2025, 2, 21)
datetime_object.strftime("%B %d, %Y")
```

---

#### Perform Arithmetic with Dates

* You'll likely need to find differences in times or perform some other arithmetic calculation
* Common arithmetic operations and `timedelta` provide you with what you need. 

```python
date_1 = datetime(2025, 2, 21, 10, 15)
date_2 = datetime(2025, 3, 6, 11, 30)
```

```python
difference = date_2 - date_1
difference
```

```python
difference.total_seconds()
```

---

## Federal Reserve Bank of St. Louis

<img src="./fred-logo-2x.png" alt="FRED logo"/>

The [Federal Reserve Bank of St. Louis](https://fred.stlouisfed.org/) (FRED) is an online database provided by the Federal Reserve Bank of St. Louis that offers access to a vast collection of economic data. It contains over 765,000 U.S. and international time series data, including key indicators like GDP, unemployment rates, inflation, interest rates, and more. FRED serves as a valuable resource for researchers, policymakers, and analysts to retrieve, analyze, and visualize economic trends and financial data. The platform provides tools to create customized charts, download data in various formats, and explore historical economic patterns, making it a critical tool for economic analysis.

---

### Consumer Price Index

The Consumer Price Index (CPI) is a measure used to track changes in the average prices of a basket of goods and services typically purchased by households over time. It is a key indicator of inflation, reflecting how much the cost of living has increased or decreased. The CPI is calculated by comparing the price of a fixed set of goods and services, which includes items like food, housing, transportation, and healthcare, from one period to another. Governments and economic analysts use the CPI to assess the purchasing power of money and to adjust policies such as interest rates and social security benefits.

#### Task 01 📍🔎

<!-- BEGIN QUESTION -->

Before you start working with real CPI data, import everything from the `datascience` library.

```python
...
```

<!-- END QUESTION -->

---

#### Task 02 📍

Create the table `cpi` from the `cpi.csv` file collected from [FRED](https://fred.stlouisfed.org/series/CPIAUCSL).

```python
cpi = ...
cpi
```

```python
grader.check("task_02")
```

---

#### Task 03 📍🔎

<!-- BEGIN QUESTION -->

What are the data types for the `'Date'` and `'CPI'` values in the `cpi` table?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

### Interpreting CPI Values

You should see that the CPI value for **January 1, 1947**, is **21.48**. What does that number mean?  

- The **Consumer Price Index (CPI)** measures the average price of a standard basket of goods over time.  
- The index is set relative to a **base period**, which in this case is **1982–1984**, where the average CPI is **100**.
    - The Bureau of Labor Statistics choose this time period for the reference because it was a relatively stable economic period without extreme inflation or deflation.
    - They periodically update the baseline when needed.
- A **CPI of 21.48 in 1947** means that, on average, prices in 1947 were **21.48% of the prices in 1982–1984**.  
- The **largest CPI in the data (319.086) from January 1, 2025** suggests that prices in early 2025 are **a little over 3 times the average price level of 1982–1984**.  


---

### Matplotlib

<img src="./matplotlib_logo_light.svg" width=300px alt="Matplotlib logo"/>

* One of the primary data visualization libraries for Python is [Matplotlib](https://matplotlib.org/)
* You are not expected to work directly with Matplotlib in MATH 108
* The `datascience` library contains table methods that allow to interact with Matplotlib indirectly.
    * e.g. The [`datascience` table method `plot`](https://datascience.readthedocs.io/en/master/_autosummary/datascience.tables.Table.plot.html#datascience.tables.Table.plot) creates a line plot with Matplotlib's function [`matplotlib.pyplot.plot`](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.plot.html).

---

#### Configuring Jupyter

Run the following command to configure his notebook to show visualizations in the same way you see them for MATH 108.

```python
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

- `%matplotlib inline` ensures that **Matplotlib plots are displayed directly inside the Jupyter Notebook**, rather than in a separate window.  
- `import matplotlib.pyplot as plt` loads Matplotlib’s plotting library.  
- `plt.style.use('fivethirtyeight')` applies the **"FiveThirtyEight" style**, which provides a visually appealing, publication-quality design for charts. This matches the style used in some course materials.  

---

### Task 04 📍🔎

<!-- BEGIN QUESTION -->

Try to create a line plot using the `plot` table method with the `cpi` table where the horizontal axis shows the dates and the vertical axis shows the CPI values. _This image should not look good by default._

```python
...

# The following adds a title to the graphic.
plt.title('CPI Over Time')
plt.show()
```

<!-- END QUESTION -->

---

### Task 05 📍

Create a function called `str_to_datetime` based on `datetime.strptime` that converts strings like those presented in the `cpi` table into `datetime` format. For this you need to know that a date string like `'1947-01-01'` is in the format `'%Y-%m-%d'`

```python
...
```

```python
grader.check("task_05")
```

---

### Task 06 📍

Covert all the date values in the `cpi` table into `datetime` format. Assign the resulting table to `cpi_for_plotting`.

```python
cpi_for_plotting = ...
cpi_for_plotting
```

```python
grader.check("task_06")
```

---

### Task 07 📍🔎

<!-- BEGIN QUESTION -->

Now, create a line plot using the `plot` table method with the `cpi` table where the horizontal axis shows the (updated) dates and the vertical axis shows the CPI values. _This image should look much better compared to the previous attempt._

```python
...

# The following adds a title to the graphic.
plt.title('CPI Over Time')
plt.show()
```

<!-- END QUESTION -->

---

### Task 08 📍🔎

<!-- BEGIN QUESTION -->

What is the general trend of CPI values over time based on this graphic?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

### Identifying Sudden Big Changes

You may notice sharp increases or decreases in CPI values over time. However, when these changes happen quickly, they can be difficult to detect in a basic line graph, but have a profound impact on the economy. Additionally, the overall trend of CPI might obscure short-term fluctuations, making it harder to see changes in relation to surrounding dates. For these reasons, analysts often calculate relative changes, such as month-to-month percentage changes, and plot those values. This approach removes the long-term trend from focus, allowing the rate of change to become the center of attention.

---

### Task 09 📍

Import numpy as `np` to access tools to help you calculate the month-to-month percentage changes in CPI values.

```python
...
```

```python
grader.check("task_09")
```

---

#### Insert a Value into an Array

The [NumPy function `np.insert`](https://numpy.org/doc/2.1/reference/generated/numpy.insert.html#numpy.insert) provides you with a way to add a value into an array at a specific location. For example, if you have defined 

`arr1 = make_array(2, 3, 4, 5)`

and you want add the value `1` to the start of the array, you can run

`np.insert(arr1, 0, 1)`

to create a new array with all 5 values and 1 at the start (index 0) of the array.

```python
arr1 = make_array(2, 3, 4, 5)
np.insert(arr1, 0, 1)
```

---

#### Delete a Value from an Array

The [NumPy function `np.delete`](https://numpy.org/doc/2.1/reference/generated/numpy.delete.html#numpy.delete) provides you with a way to remove a value from an array at a specific location. For example, if you have defined 

`arr2 = make_array(1, 2, 3, 4, 5)`

and you want remove the value `3` at the 3rd position of the array, you can run

`np.delete(arr2, 2)`

to create a new array with the remaining 4 values.

```python
arr2 = make_array(1, 2, 3, 4, 5)
np.delete(arr2, 2)
```

---

### Task 10 📍

Create a table like `cpi_for_plotting` called `cpi_percent_change` with an additional column `'CPI % Change (Monthly)'` at the right-end of the table showing the positive or negative change in CPI values from month to month as a percentage.

For example, from the relative percentage change from the CPI value of 21.48 to 21.62 is:

$$\frac{21.62 - 21.48}{21.48} * 100 \approx 0.65\%$$

Showing that the CPI change from January to February in 1947 was a little over one half percent.

```python
...
```

```python
grader.check("task_10")
```

---

### Task 11 📍🔎

<!-- BEGIN QUESTION -->

Use `cpi_percent_change` to plot the month-to-monthly relative percentage changes in CPI values for the dates starting in the year 2000.

```python
...

# The following adds a title to the graphic.
plt.title('Month-to-Month CPI % Change')
plt.show()
```

<!-- END QUESTION -->

---

### Recessions

From [Investopedia](https://www.investopedia.com/terms/r/recession.asp)

> A recession is a significant, widespread, and prolonged downturn in economic activity. 

* The most recent recession in the United States was from February to April 2020, during the COVID-19 pandemic.
* The Great Recession lasted from December 2007 to June 2009, and was tied to the global financial crisis.

Though CPI is meant to measure inflation and is not used to define if there is a recession, a significant change in CPI can be a warning sign of a recession. 

---

### Task 12 📍🔎

<!-- BEGIN QUESTION -->

What evidence in the visualization of month-to-month CPI percentage changes do you see for the two most recent US recessions?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Reflection

In this Guided Learning Activity, you explored how to visualize trends over time, focusing on working with date and time data. You learned how to convert date strings into `datetime` format, sort time-based data, and create visualizations to analyze trends. By applying these techniques to real-world data (such as CPI values from FRED), you practiced identifying patterns, anomalies, and economic events through data visualization. The activity reinforced the importance of relative percentage changes for understanding short-term fluctuations, rather than relying solely on absolute values.

---

## License

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a>.

<img src="./by-nc-sa.png" width=100px>

Tests for task_02 (points: None)
```
>>> isinstance(cpi, Table)
True
```
failure_message: "❌ cpi should be a Table."
success_message: "✅ cpi is a Table."

```
>>> cpi.num_rows
937
```
failure_message: "❌ cpi does not have the correct number of rows."
success_message: "✅ cpi has the correct number of rows."

```
>>> cpi.num_columns
2
```
failure_message: "❌ cpi does not have the correct number of columns."
success_message: "✅ cpi has the correct number of columns."

```
>>> cpi
Date       | CPI
1947-01-01 | 21.48
1947-02-01 | 21.62
1947-03-01 | 22
1947-04-01 | 22
1947-05-01 | 21.95
1947-06-01 | 22.08
1947-07-01 | 22.23
1947-08-01 | 22.4
1947-09-01 | 22.84
1947-10-01 | 22.91
... (927 rows omitted)
```
failure_message: "❌ cpi does not seem correct."
success_message: "✅ cpi seems correct."



Tests for task_05 (points: None)
```
>>> callable(str_to_datetime)
True
```
failure_message: "❌ str_to_datetime is not callable"
success_message: "✅ str_to_datetime is callable."

```
>>> isinstance(str_to_datetime('1947-01-01'), datetime)
True
```
failure_message: "❌ str_to_datetime does not correctly converts the date string '1947-01-01' to datetime format."
success_message: "✅ str_to_datetime correctly converts the date string '1947-01-01' to datetime format."

```
>>> str_to_datetime('1947-01-01')
datetime.datetime(1947, 1, 1, 0, 0)
```
failure_message: "❌ str_to_datetime does not seem to be working correctly."
success_message: "✅ str_to_datetime seems to be working correctly."



Tests for task_06 (points: None)
```
>>> isinstance(cpi_for_plotting, Table)
True
```
failure_message: "❌ cpi_for_plotting should be a Table."
success_message: "✅ cpi_for_plotting is a Table."

```
>>> cpi_for_plotting.num_rows
937
```
failure_message: "❌ cpi_for_plotting does not have the correct number of rows."
success_message: "✅ cpi_for_plotting has the correct number of rows."

```
>>> cpi_for_plotting.num_columns
2
```
failure_message: "❌ cpi_for_plotting does not have the correct number of columns."
success_message: "✅ cpi_for_plotting has the correct number of columns."

```
>>> cpi_for_plotting.labels
('Date', 'CPI')
```
failure_message: "❌ cpi_for_plotting does not have the correct labels."
success_message: "✅ cpi_for_plotting has the correct labels."

```
>>> isinstance(cpi_for_plotting.column('Date').item(0), datetime)
True
```
failure_message: "❌ The Date values are not in datetime format."
success_message: "✅ The Date values are in datetime format."

```
>>> isinstance(cpi_for_plotting.column('CPI').item(0), float)
True
```
failure_message: "❌ The cpi_for_plotting values are not in float format."
success_message: "✅ The cpi_for_plotting values are in float format."



Tests for task_09 (points: None)
```
>>> isinstance(np.array([]), np.ndarray)
True
```
failure_message: ❌ NumPy has not be imported as np.
success_message: ✅ NumPy has been imported as np.


Tests for task_10 (points: None)
```
>>> isinstance(cpi_percent_change, Table)
True
```
failure_message: "❌ cpi_percent_change should be a Table."
success_message: "✅ cpi_percent_change is a Table."

```
>>> cpi_percent_change.num_rows
937
```
failure_message: "❌ cpi_percent_change does not have the correct number of rows."
success_message: "✅ cpi_percent_change has the correct number of rows."

```
>>> cpi_percent_change.num_columns
3
```
failure_message: "❌ cpi_percent_change does not have the correct number of columns"
success_message: "✅ cpi_percent_change has the correct number of columns."

```
>>> cpi_percent_change.labels
('Date', 'CPI', 'CPI % Change (Monthly)')
```
failure_message: "❌ cpi_percent_change does not have the correct labels."
success_message: "✅ cpi_percent_change has the correct labels."

```
>>> isinstance(cpi_percent_change.column('Date').item(0), datetime)
True
```
failure_message: "❌ cpi_percent_change does not have the Date values in datetime format."
success_message: "✅ cpi_percent_change has the Date values in datetime format."

```
>>> cpi_percent_change.column('CPI % Change (Monthly)').item(0)
0.0
```
failure_message: "❌ The first % change is not correct."
success_message: "✅ The first % change is correct."

```
>>> np.isclose(cpi_percent_change.column('CPI % Change (Monthly)').item(1), 0.6517690875232801)
True
```
failure_message: "❌ The % change from January to February in 1947 is not correct."
success_message: "✅ The % change from January to February in 1947 is correct."

```
>>> np.sum(cpi_percent_change.column('CPI % Change (Monthly)'))
270.76862523870034
```
failure_message: "❌ cpi_percent_change does not seem correct."
success_message: "✅ cpi_percent_change seems to have the correct ."

