```python
# Initialize Otter
import otter
grader = otter.Notebook("gla07.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Guided Learning Activity 07: Simulation-Based Hypothesis Testing

This Guided Learning Activity is designed for you to complete alongside a Data Ambassador from the course. You might find that it feels like a combination of the lectures and lab assignment. Whether you are participating live or watching the recording of the live meeting, let the Data Ambassador guide you through the following tasks. There will be moments for you to reflect and explore your own ideas as a way to solidify concepts and skills introduced by your instructor. Keep in mind that this is not a graded assignment for MATH 108 by default. If you have any concerns about participation, reach out to your instructor.

---

## Learning Objectives

1. Outline the general process of a simulation-based hypothesis test.
2. Perform operations on tables and arrays.
3. Visualize distributions.
4. Iteratively simulate generating data based on some hypothesis.
5. Perform a simulation-based hypothesis test.
6. Reflect on the possible error of the conclusion from a hypothesis test.

---

## Configure the Notebook

Run the following code cell to set up the notebook.

```python
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

---

## Hypothesis Testing

Hypothesis testing was [initially developed](https://en.wikipedia.org/wiki/Statistical_hypothesis_test) as a tool for guiding research.

```mermaid
graph LR;
    A["Research under Null Hypothesis"] --> B["Propose Alternative Hypothesis"]
    B --> C{"Does evidence support \n rejecting the Null Hypothesis?"}
    C -->|"No (Fail to reject Null)"| E["Continue research under Null Hypothesis"]
    C -->|"Yes (Reject Null)"| D["Shift research towards Alternative Hypothesis"]
```

* Provides a framework for rejecting claims surrounding a random process.
* Decisions are based on empirical evidence.
* Considers that the results of the decision could be wrong.

---

### Statistical Significance

How is the decision to reject the null hypothesis made (Simulation-Based Approach)?

* **Observe data:** Collect the empirical data relevant to the research question.
* **Define a testable hypothesis:** Formulate a null hypothesis and an alternative hypothesis based on the data.
* **Develop a test statistic:** Choose a statistic (e.g., mean difference, correlation) that quantifies the effect of interest.
* **Apply the test statistic to the observed data:** Calculate the test statistic value using the collected data.
* **Derive a distribution for the test statistic under the null hypothesis (Simulation):**
    * Develop a function to randomly generate data assuming the null hypothesis is true.
    * Run the function many times (e.g., thousands of simulations).
    * Calculate the test statistic for each simulated dataset.
    * Collect the test statistic values from all simulations to create a distribution.
* **Compare the observed test statistic to the simulated distribution:** Determine the position of the observed test statistic within the simulated distribution.
* **Calculate the p-value:** Determine the proportion of simulated test statistic values that are equal to or more extreme than the observed test statistic value.
* **Make a decision:**
    * If the p-value is below the pre-defined significance level (alpha), reject the null hypothesis.
    * If the p-value is greater than or equal to alpha, fail to reject the null hypothesis.

---

## Is Yawning Contagious?

<a href="https://unsplash.com/photos/yawning-white-and-brown-short-coated-puppy-IE2Z11zKsso"><img src="./daniel-lincoln-IE2Z11zKsso-unsplash.jpg" width=300px alt="yawning white and brown short coated puppy"></a>

This activity follows the research from [_Interspecific Contagious Yawning in Humans_](https://pmc.ncbi.nlm.nih.gov/articles/PMC9332820/) by Andrew C Gallup and Sabina Wozny published in 2022.

---

### Research Summary

* Contagious yawning is well-documented in humans and social vertebrates, with some non-human animals yawning in response to human caregivers.
* This study investigated whether humans yawn contagiously in response to non-human animals and if this is influenced by phylogenetic relatedness or domestication.
* Researchers analyzed self-reported yawning behavior from 296 participants exposed to yawning stimuli from various animal groups, including fish, amphibians, reptiles, birds, mammals, and domesticated pets.
* The study found strong evidence for interspecific yawn contagion, meaning people yawned in response to animal yawns.
* However, yawning frequency and likelihood did not vary significantly across different animal groups, suggesting that diverse yawning stimuli can trigger contagious yawning regardless of species.

---

## Participants

* The researchers collected data from 167 men and 137 women using Amazon Mechanical Turk (MTurk).
* MTurk is a crowdsourcing platform that allows businesses, researchers, and individuals to outsource small tasks, known as Human Intelligence Tasks (HITs), to a large pool of remote workers.
* The participants needed to:
    * Claim residency status in the United States
    * Have a successful completion rate > 95%
    * Completed a minimum of 50 tasks
    * Pass a few checks to distinguish between computers and humans

---

## Study Design

* Participants were shown 12 pairs of images and were asked to pick (left or right) based on which image depicted the subject yawning
```mermaid
graph LR;
    A["Participants"] --> B["Random Assignment to One Condition"]
    B --> C["Building Windows (Control)"]
    B --> D["Fish"]
    B --> E["Amphibians"]
    B --> F["Reptiles"]
    B --> G["Birds"]
    B --> H["Non-Primate Mammals"]
    B --> I["Apes"]
    B --> J["Pets"]

    C --> K["12 Open + 12 Closed Windows"]
    D --> L["12 Yawning + 12 Non-Yawning Fish"]
    E --> M["12 Yawning + 12 Non-Yawning Amphibians (Frogs, Salamanders)"]
    F --> N["12 Yawning + 12 Non-Yawning Reptiles (Monitors, Tortoises, Geckos, etc.)"]
    G --> O["12 Yawning + 12 Non-Yawning Birds (Shorebirds, Birds of Prey, etc.)"]
    H --> P["12 Yawning + 12 Non-Yawning Mammals (Elephant, Camel, Panda, etc.)"]
    I --> Q["12 Yawning + 12 Non-Yawning Apes (Chimpanzees, Gorillas, etc.)"]
    J --> R["12 Yawning + 12 Non-Yawning Pets (Dogs, Cats)"]
```
* Afterwards, participants were then asked
    *  Whether they yawned while reviewing the stimuli (yes/no), and if so, how many times
    *  How many hours of sleep they had the previous night
    *  How tired they were on a 10-point scale (1: not tired at all; 10: extremely tired)

```mermaid
graph TD;
    A["Post-Viewing Questions"] --> B["Most Satisfying Image (Not Analyzed)"]
    A --> C{"Did You Yawn? (Yes/No)"}
    C --> D["How Many Times?"]
    A --> E["Sleep Hours Last Night"]
    A --> F["Tiredness Rating (1-10)"]
```

---

## The Data

The datasets used to generate the results in the paper are provided in `Comparative CY dataset FINAL.csv` and sourced from [Harvard Dataverse](https://doi.org/10.7910/DVN/3STP5S).

### Task 01 📍

Create the table `results` from the CSV file `Comparative CY dataset FINAL.csv`.

```python
results = ...
results
```

```python
grader.check("task_01")
```

---

## Analyzing Yawn Frequency

### Task 02 📍

In the analysis portion of the paper, the researchers write:

> The complete distribution of yawning frequency revealed eight outliers, as defined by 1.5 times the interquartile range. In each case, participants reported more than 12 contagious yawns during the stimulus evaluation (which took ~1 min to complete in pretesting), or greater than 1 yawn per stimulus pairing. After excluding these participants from the analysis, the final sample was 296 (control N = 38; fish N = 32; amphibians N = 38; reptiles N = 34; birds N = 34; non-primate mammals N = 37; apes N = 44; pets N = 39).

Confirm these counts for the eight groups by creating a table called `by_condition` that contains two columns `'Condition'` and `'count'` where each row shows the frequency of the condition in the `results` table.

```python
by_condition = ...
by_condition
```

```python
grader.check("task_02")
```

---

### Task 03 📍🔎

<!-- BEGIN QUESTION -->

How are the yawn frequencies distributed among each group? Create a histogram to visualize the distribution of yawn frequency values for each condition. We recommend using `overlay=False` with the [`hist` method](https://datascience.readthedocs.io/en/master/_autosummary/datascience.tables.Table.hist.html#datascience.tables.Table.hist). Then, provide a brief summary of your observations. (Do any conditions stand out in their distribution?)

_Type your answer here, replacing this text._

```python
...
```

<!-- END QUESTION -->

---

### Task 04 📍

How do the yawn frequencies compare across the eight conditions? Create a table called `ave_yawn_by_condition` that has two columns:
* `'Condition'`: the eight condition labels as presented in `results`.
* `'Average Yawn Frequency'`: The average yawn frequency for the participants exposed to each of the eight conditions.

```python
ave_yawn_by_condition = ...
ave_yawn_by_condition
```

```python
grader.check("task_04")
```

---

### Lowest Average Frequency

Run the following code cell to generate a an interactive bar chart that allows you to compare the average yawn frequencies between the condition groups.

```python
ave_yawn_by_condition.ibar('Condition', 
                           title='Average Yawn Frequency by Condition')
```

Examine the graphic. Notice that the averages do not seem all equal. The lowest average frequency (approx. 2.08 yawns) is for the control group (windows). This suggests that all the other conditions have some measurable effect on human yawning, where birds might have the strongest effect and amphibians the weakest effect.

---

### Forming Hypotheses

### Task 06 📍🔎

<!-- BEGIN QUESTION -->

Form two hypotheses to assess whether observing yawning in a non-human animal has a significant effect on human yawning based on average yawn frequencies.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Testing Against the Null Hypothesis

Now, you can implement a simulation-based hypothesis test to measure the evidence against the null hypothesis. 

---

### Test Statistic

For a test statistic, you will use the average yawn frequency for the control group. What you can allow to vary is the way the participants were assigned to the conditions. Maybe it was just by chance that this particular group of participants assigned to the control group had a low average yawn frequency.

---

### Task 07 📍

What is the observed value of the test statistic? Using the data in `results`, assign the average yawn frequency (`float`) for the control group to `observed_average`.

```python
observed_average = ...
observed_average
```

```python
grader.check("task_07")
```

---

### Task 08 📍

Create a function called `simulate_avg_yawn_frequency_control` that takes no input, generates a random sample from `results` of the same size as the original control group and returns the average yawn frequency (`float`) for that random sample.

```python
results
```

```python
# Define the function
def simulate_avg_yawn_frequency_control():
    ...

# Call the function
simulate_avg_yawn_frequency_control()
```

```python
grader.check("task_08")
```

---

### Task 09 📍

Form 10,000 iterations of `simulate_avg_yawn_frequency_control()` and save the resulting test statistics in the array `simulated_average_yawn_frequencies`. 

**Note:** This code will take about a minute to run.

```python
simulated_average_yawn_frequencies = ...
...
```

```python
grader.check("task_09")
```

---

### Distribution of Test Statistic

### Task 10 📍🔎

<!-- BEGIN QUESTION -->

Visualize the distribution of test statistic values based on the 10,000 simulations. 

**Note**: We've provided code to add a dot on the histogram indicating the average yawn frequency for the observed control group.

```python
...

# Add a point for the observed value of the test statistic (observed_average)
plt.ylim(-0.1, 1)
plt.scatter(observed_average, 0, color='red', s=30, label='Observed Test Statistic');
plt.title('Distribution of Test Statistics')
plt.legend()
plt.show()
```

<!-- END QUESTION -->

---

### Task 11 📍🔎

<!-- BEGIN QUESTION -->

Calculate the p-value based on the results in `simulated_average_yawn_frequencies`. From your p-value, what can you conclude from this test if you use a standard 5% p-value cutoff?

_Type your answer here, replacing this text._

```python
p_value = ...
p_value
```

<!-- END QUESTION -->

---

### Task 12 📍🔎

<!-- BEGIN QUESTION -->

What type of error (Type I or II) could you have made based on your conclusion? If possible, what is the chance that you made that error?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Reflection

In this activity, you practiced going through a simulation-based hypothesis test based on a real research paper. You used computational tools to perform calculations, visualize distributions, simulate generating data, run many simulations, and calculate a p-value. Lastly, you considered that the results of your test may not be true.

---

## License

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a>.

<img src="./by-nc-sa.png" width=100px>

Tests for task_01 (points: None)
```
>>> isinstance(results, Table)
True
```
failure_message: "❌ results should be assigned to a Table."
success_message: "✅ results is assigned to a Table."

```
>>> results.num_rows
296
```
failure_message: "❌ results does not have the correct number of rows."
success_message: "✅ results has the correct number of rows."

```
>>> results.labels == ('Condition', 'Age (Years)', 'Gender (1:Male; 2: Female)', 'Yawned (1:yes; 0: no)', 'Yawn Frequency', 'Sleep (Hours)', 'Tiredness Scale (1-10)')
True
```
failure_message: "❌ results does not have the correct column labels."
success_message: "✅ results has the correct column labels."

```
>>> results
Condition | Age (Years) | Gender (1:Male; 2: Female) | Yawned (1:yes; 0: no) | Yawn Frequency | Sleep (Hours) | Tiredness Scale (1-10)
0.Control | 26          | 1                          | 1                     | 2              | 8             | 8
0.Control | 28          | 2                          | 0                     | 0              | 8             | 2
0.Control | 55          | 2                          | 0                     | 0              | 7             | 2
0.Control | 34          | 2                          | 0                     | 0              | 7             | 4
0.Control | 32          | 1                          | 0                     | 0              | 7             | 1
0.Control | 30          | 1                          | 1                     | 9              | 8             | 10
0.Control | 30          | 1                          | 1                     | 12             | 5             | 9
0.Control | 36          | 2                          | 0                     | 0              | 8             | 3
0.Control | 35          | 1                          | 0                     | 0              | 8             | 10
0.Control | 31          | 2                          | 1                     | 5              | 8             | 8
... (286 rows omitted)
```
failure_message: "❌ results does not seem correct."
success_message: "✅ results seems correct."



Tests for task_02 (points: None)
```
>>> isinstance(by_condition, Table)
True
```
failure_message: "❌ by_condition should be assigned to a Table."
success_message: "✅ by_condition is assigned to a Table."

```
>>> by_condition.num_rows
8
```
failure_message: "❌ by_condition does not have the correct number of rows."
success_message: "✅ by_condition has the correct number of rows."

```
>>> by_condition.labels == ('Condition', 'count')
True
```
failure_message: "❌ by_condition does not have the correct column labels."
success_message: "✅ by_condition has the correct column labels."

```
>>> by_condition
Condition | count
0.Control | 38
1.Fish    | 32
2.Amphib  | 38
3.Reptile | 34
4.Bird    | 34
5.Mammal  | 37
6.Ape     | 44
7.Pet     | 39
```
failure_message: "❌ by_condition does not seem correct."
success_message: "✅ by_condition seems correct."



Tests for task_04 (points: None)
```
>>> isinstance(ave_yawn_by_condition, Table)
True
```
failure_message: "❌ ave_yawn_by_condition should be assigned to a Table."
success_message: "✅ ave_yawn_by_condition is assigned to a Table."

```
>>> ave_yawn_by_condition.num_rows
8
```
failure_message: "❌ ave_yawn_by_condition does not have the correct number of rows."
success_message: "✅ ave_yawn_by_condition has the correct number of rows."

```
>>> ave_yawn_by_condition.labels == ('Condition', 'Average Yawn Frequency')
True
```
failure_message: "❌ ave_yawn_by_condition does not have the correct column labels."
success_message: "✅ ave_yawn_by_condition has the correct column labels."

```
>>> ave_yawn_by_condition
Condition | Average Yawn Frequency
0.Control | 2.07895
1.Fish    | 3.375
2.Amphib  | 2.55263
3.Reptile | 3.58824
4.Bird    | 4.02941
5.Mammal  | 2.62162
6.Ape     | 2.93182
7.Pet     | 3.4359
```
failure_message: "❌ ave_yawn_by_condition does not seem correct."
success_message: "✅ ave_yawn_by_condition seems correct."



Tests for task_07 (points: None)
```
>>> import numbers
>>> isinstance(observed_average, numbers.Number)
True
```
failure_message: "❌ observed_average should be assigned to a numerical value."
success_message: "✅ observed_average is assigned to a numerical value."

```
>>> np.isclose(observed_average, 2.0789473684210527)
True
```
failure_message: "❌ observed_average does not seem correct."
success_message: "✅ observed_average seems correct."



Tests for task_08 (points: None)
```
>>> callable(simulate_avg_yawn_frequency_control)
True
```
failure_message: "❌ simulate_avg_yawn_frequency_control is not callable."
success_message: "✅ simulate_avg_yawn_frequency_control is callable."

```
>>> isinstance(simulate_avg_yawn_frequency_control(), float) and simulate_avg_yawn_frequency_control() >= 0
True
```
failure_message: "❌ simulate_avg_yawn_frequency_control returns a non-negative float value."
success_message: "✅ simulate_avg_yawn_frequency_control returns a non-negative float value."



Tests for task_09 (points: None)
```
>>> isinstance(simulated_average_yawn_frequencies, np.ndarray)
True
```
failure_message: "❌ simulated_average_yawn_frequencies is not an array."
success_message: "✅ simulated_average_yawn_frequencies is an array."

```
>>> len(simulated_average_yawn_frequencies) == 10000
True
```
failure_message: "❌ simulated_average_yawn_frequencies should contain 10,000 items."
success_message: "✅ simulated_average_yawn_frequencies contains 10,000 items."

