```python
# Initialize Otter
import otter
grader = otter.Notebook("gla03.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Guided Learning Activity 03: Accessing Data

This Guided Learning Activity is designed for you to complete alongside a Data Ambassador from the course. You might find that it feels like a combination of the lectures and lab assignment. Whether you are participating live or watching the recording of the live meeting, let the Data Ambassador guide you through the following tasks. There will be moments for you to reflect and explore your own ideas as a way to solidify concepts and skills introduced by your instructor. Keep in mind that this is not a graded assignment for MATH 108 by default. If you have any concerns about participation, reach out to your instructor.

---

## Learning Objectives

1. Recognize different data storage formats, including `CSV` and `JSON`.
2. Learn how to load and explore datasets using the `datascience` library.
3. Practice accessing and extracting specific data from tables using sorting and filtering techniques.
4. Map data that contains latitude and longitude information.
5. Utilize APIs to access information directly from remote sources.

---

## Accessing Stored Data

In data science, working with real-world data often means dealing with a variety of data sources, from structured files like `CSV`s and JSONs to dynamic sources accessed via APIs. Understanding how to read, process, and analyze data from these different formats is a crucial skill for any data scientist. Whether you're working with a static dataset stored locally or pulling live data from a web service, the ability to efficiently load and manipulate data enables deeper insights and more informed decision-making.

---

### Data Formats

<a href="https://www.flaticon.com/packs/file-type-collection"><img src="./csv_icon.png" alt="CSV file icon" width=100px><img src="./json_icon.png" alt="JSON file icon" width=100px><img src="./txt_icon.png" alt="TXT file icon" width=100px><img src="./xls_icon.png" alt="XLS file icon" width=100px><img src="./xml_icon.png" alt="XML file icon" width=100px></a> ...

Data can be stored in various file formats, each designed for different use cases and levels of complexity. Although there are a large number of possible file formats, you'll likely run into some common file formats such as XLSX, CSV, and `JSON`. You can usually recognize a file's format by the extension on the file name. XLSX files end with `.xlsx`, CSV files end with `.csv`, and `JSON` files end with `.json` for example. Each file type should come with good documentation and even some tutorials to help you better understand how to extract the information you need.

---

#### CSV

A `CSV` (Comma Separated Values) file is a simple text file that stores tabular data (numbers and text) in plain text. Each line of the file represents a row of data, and the values within each row are separated by commas. You primarily work indirectly with `CSV` files in MATH 108. The following example from [Wikipedia's Comma-separated values page](https://en.wikipedia.org/wiki/Comma-separated_values) showcases the general format of a `CSV` file:

```csv
Year,Make,Model,Description,Price
1997,Ford,E350,"ac, abs, moon",3000.00
1999,Chevy,"Venture ""Extended Edition""","",4900.00
1999,Chevy,"Venture ""Extended Edition, Very Large""","",5000.00
1996,Jeep,Grand Cherokee,"MUST SELL!
air, moon roof, loaded",4799.00
```
The first line shows the variable labels for each car. In general, each row provides the details for each type of car. Since the data is organized by commas, not by rows, the fact that the Grand Cherokee is split into two lines is no problem.

---

#### JSON

JSON (JavaScript Object Notation) files are a flexible, lightweight format often used for structured data as they represent data as key-value pairs (`"key": "value"`). The following example from [W3Schools's JavaScript `JSON`
 page](https://www.w3schools.com/js/js_json.asp) showcases the `JSON` structure:

```json
{
    "employees":
    [
        {
            "firstName": "John", 
            "lastName": "Doe"
        },
        {
            "firstName": "Anna", 
            "lastName": "Smith"
        },
        {
            "firstName": "Peter", 
            "lastName": "Jones"
        }
    ]
}
```
You don't work with these files in MATH 108, but they are common enough that you should be aware of their existence. 

---

## `datascience` Library

In MATH 108, you utilize the Python library `datascience` developed by UC Berkeley to complete most of the coursework. Run the following code cell to import all of the `datascience` library.

```python
from datascience import *
```

---

## Tables

<a href="https://www.nssl.noaa.gov/education/svrwx101/hail/types/"><img src="./hail.jpg" alt="Hail examples" width=800px></a>

NOAA has [data reports from severe weather incidents](https://www.ncei.noaa.gov/access/metadata/landing-page/bin/iso?id=gov.noaa.ncdc:C00773) including the size of the hail and the time and location the data was collected. It could be pretty challenging to analyze that data if the size, date, and location information were stored separately. A table is a way to organize this information neatly, with each item (hail report) getting its own row, and each piece of information (like size and location) getting its own column. This makes it much easier to find and compare information than navigating the arrays/lists separately.  The `datascience` library in Python provides a data type called `Table` for this purpose. It's a powerful way to keep your information structured and easy to work with, giving you tools to add, find, and analyze your data.

<a href="https://cloud.google.com/datasets"><img src="./NOAA_hail_reports.png" alt="Data from NOAA on severe weather hail reports." width=600px></a>

---

### `read_table`

The function [`read_table`](https://datascience.readthedocs.io/en/master/_autosummary/datascience.tables.Table.read_table.html#datascience.tables.Table.read_table) in the `Table` module provides you with a way to load data from `CSV` files as a `Table` data type. Here is the documentation:

```python
Signature: Table.read_table(filepath_or_buffer, *args, **vargs)
Docstring:
Read a table from a file or web address.

Args:
    filepath_or_buffer -- string or file handle / StringIO; The string
                      could be a URL. Valid URL schemes include http,
                      ftp, s3, and file.

Returns:
    a table read from argument
                      
Example:

>>> Table.read_table('https://www.inferentialthinking.com/data/sat2014.csv')
State        | Participation Rate | Critical Reading | Math | Writing | Combined
North Dakota | 2.3                | 612              | 620  | 584     | 1816
Illinois     | 4.6                | 599              | 616  | 587     | 1802
Iowa         | 3.1                | 605              | 611  | 578     | 1794
South Dakota | 2.9                | 604              | 609  | 579     | 1792
Minnesota    | 5.9                | 598              | 610  | 578     | 1786
Michigan     | 3.8                | 593              | 610  | 581     | 1784
Wisconsin    | 3.9                | 596              | 608  | 578     | 1782
Missouri     | 4.2                | 595              | 597  | 579     | 1771
Wyoming      | 3.3                | 590              | 599  | 573     | 1762
Kansas       | 5.3                | 591              | 596  | 566     | 1753
... (41 rows omitted)
        
File:      /srv/conda/envs/notebook/lib/python3.11/site-packages/datascience/tables.py
Type:      method
```

You essentially just need to provide the function with the location of the `CSV` file as input and it will produce the table for you. As long as the `CSV` file is formatted correctly, you should end up with a decent table.

---

#### Task 01 💻

Use `Table.read_table` to assign `hail` to a table containing the contents of the `CSV` file `hail.csv` that we've put into your JupyterHub folder for this activity.

```python
hail = ...
hail
```

```python
grader.check("task_01")
```

---

### `Table` Documentation

Running the following command should reveal that `hail` is a `Table`.

```python
type(hail)
```

The `Table` data type comes with several attributes and methods for you to work with. You can reference the [`datascience` documentation page on Tables](https://datascience.readthedocs.io/en/master/tables.html) to get more details about working with tables. 

The following tasks are meant to give you some practice working with this data type.

---

#### Task 02 💻

How many reports are there in the `hail` table? Assign this number to `hail_num_rows`.

```python
hail_num_rows = ...
hail_num_rows
```

```python
grader.check("task_02")
```

---

#### Task 03 💻

Calculate the average hail size from the `hail` table and assign that value to `average_hail_size`. We've imported NumPy for you on the first line of the following code cell.

**Note**: The numbers in the `'size'` column of `hail` are expressed in 1/100 of an inch, so 175 = 1.75 inches.

```python
import numpy as np
average_hail_size = ...
average_hail_size
```

```python
grader.check("task_03")
```

---

#### Task 04 💻

Create a new table called `KS_hail` which contains the same information in `hail` filtered to just those reports from Kansas. Since `KS_hail` will only contain Kansas reports, you can remove the column `state` as it is not needed anymore.

```python
KS_hail = ...
KS_hail
```

```python
grader.check("task_04")
```

---

#### Task 05 💻

Sort the `KS_hail` table so that the reports with the largest hail size are at the top and the smallest are at the bottom of the table. Name the resulting table as `KS_hail_sorted`.

```python
KS_hail_sorted = ...
KS_hail_sorted
```

```python
grader.check("task_05")
```

---

#### Task 06 💻

What is the latitude and longitude of the oldest Texas hail report in the `hail` table? Assign the latitude to `TX_lat` and the longitude to `TX_long`.

**Note**: We've typed `TX_lat, TX_long` on the last line of the code cell so that the latitude and longitude values you assign are displayed.

```python
TX_lat = ...
TX_long = ...
TX_lat, TX_long
```

```python
grader.check("task_06")
```

---

## Mapping Data

<a href="https://python-visualization.github.io/folium/latest/getting_started.html#Choropleth-maps"><img src="./USA_map.png" alt="A map of the USA showing unemployment rates by state." width=500px></a>

The `datascience` comes with a `Maps` module to help you visualize your tabular data. The tools in this module are derived from the Python library called [Folium](https://python-visualization.github.io/folium/latest/). For example, you can visualize the hail data by putting a circle on the latitude and longitude posted for each report with the following code:

```python
hail_for_map = (
    hail.select('latitude', 'longitude', 'location')
    .relabeled('latitude', 'lat')
    .relabeled('longitude', 'lon')
    .relabeled('location', 'labels')
)
Circle.map_table(hail_for_map)
```

Notice that for this work, we needed to organize the table used for making the map so that the column labels were a specific way that way `Circle` knew which values were the latitude values, longitude values, and labels for the pop ups when you click on circle marker.

```python
hail_for_map
```

You can learn more about mapping data in the [Maps page of the `datascience` documentation](https://datascience.readthedocs.io/en/master/maps.html).

---

#### Task 07 💭

<!-- BEGIN QUESTION -->

Create a circle map of the hail report data for only the reports in Texas and adjust the size of the circles based on the size of the hail indicated in the report.

**Note**: There is no auto-grader test for this task.

```python

```

<!-- END QUESTION -->

---

## Rent Board Housing

<img src="./San_Francisco.avif" alt="" width=500px>

Starting in 2022, the owners of residential housing units in San Francisco are required by law to report certain information about their units to the San Francisco Rent Board on an annual basis. The collection of this information is called [the Rent Board Housing Inventory](https://www.sf.gov/rent-board-housing-inventory). The Housing Inventory data is stored on the [DataSF Rent Board Housing Inventory page](https://data.sfgov.org/Housing-and-Buildings/Rent-Board-Housing-Inventory/gdc7-dmcn/about_data) in a few formats. The first few lines of the `CSV` file look like:

```csv
unique_id,block_num,unit_count,case_type_name,submission_year,block_address,occupancy_type,occupancy_or_vacancy_date,occupancy_or_vacancy_date_year,bedroom_count,bathroom_count,square_footage,monthly_rent,base_rent_includes_water_sewer,base_rent_includes_natural_gas,base_rent_includes_electricity,base_rent_includes_refuse_recycling,base_rent_includes_other_utilities,past_occupancy,vacancy_date,signature_date,occupancy_or_vacancy_date_history,year_property_built,point,analysis_neighborhood,supervisor_district,data_as_of,data_loaded_at
4954258646076502265,7331,219,Housing Inventory - Unit information (2023),2023,0 Block of CHUMASERO DR,Vacant,nan,2023,One-Bedroom,One bathroom,751-1000 Sq.Ft,nan,N,N,N,N,nan,No,nan,2023-03-27T00:00:00.000,nan,,POINT (-122.473318473 37.714384926),Lakeshore,7,2025-02-12T01:30:26.000,2025-02-12T06:05:50.832
5465711111802241417,7331,219,Housing Inventory - Unit information (2023),2023,0 Block of CHUMASERO DR,Vacant,nan,2023,Two-Bedroom,Two bathrooms,1001-1250 
```

How can we get that data into our notebook?
* There is an option on DataSF's page that allows us to download the data as `CSV` and then we can upload it to our JupyterHub.
* You can access the data directly through a special URL.

---

#### Task 08 💻

<img src="./DataSF_Export.png" alt="The export pop-up window on the DataSF page.">

Download the data from [DataSF Rent Board Housing Inventory page](https://data.sfgov.org/Housing-and-Buildings/Rent-Board-Housing-Inventory/gdc7-dmcn/about_data) by clicking the Export button, upload it to the same place this notebook is located, and assign the data to the table `rent_upload`.

```python
rent_upload = ...
rent_upload
```

```python
grader.check("task_08")
```

---

### APIs

Now, you should have been able to download the `CSV` file from DataSF, upload it here, and load the notebook. A more systematic approach is to use an API to obtain that information directly from DataSF. 

An API (Application Programming Interface) is like a waiter at a restaurant. The kitchen (another program or service) has all the food (data), but you can't just barge in and grab it. You need to order through the waiter (API). The waiter takes your order (your request for data), goes to the kitchen, gets the food, and brings it back to you in a nice, organized way.

```mermaid
sequenceDiagram
    participant User as You (Hungry Person)
    participant API as API (Waiter)
    participant Server as Data Server (Kitchen)

    User->>API: Request Data (e.g., rental info for Lakeshore, 2024)
    API->>Server: Fetch requested data
    Server-->>API: Return filtered data
    API-->>User: Deliver the data in a structured format

    note over User,API: API acts as a middleman, handling data requests
    note over API,Server: API retrieves only the required data
```

So, an API is a set of rules that lets different programs talk to each other. It defines what kinds of "orders" (data requests) you can make and what kind of "food" (data) you'll get back. This is super useful because it means you don't have to build everything from scratch. You can just ask another program for the information you need. For example, the `CSV` file you uploaded earlier originally came from the large `CSV` file stored at `https://data.sfgov.org/resource/gdc7-dmcn.csv`. That file is like a big menu and takes up a lot of storage space. Instead of downloading the *whole* menu, which might be huge, we can use an API to order just the dishes (data) we want.

DataSF has an [API](https://dev.socrata.com/foundry/data.sfgov.org/gdc7-dmcn) that lets us ask for specific parts of the data. We can add things to the URL, like `?submission_year=2024&analysis_neighborhood='Lakeshore'`, to get only the data we're interested in, like the 2024 data for the Lakeshore neighborhood. It's like telling the waiter, "I want the soup, but only the tomato soup, and only a small bowl." The API helps us get exactly what we need without having to deal with the whole kitchen (the entire dataset).

Run the following command to get the 2024 Lakeshore data. Notice that we broke up the long string for the URL twice using the `\` symbol. This is helpful so that string doesn't extend too far to the right of the screen.

```python
Lakeshore_2024_url = "https://data.sfgov.org/resource/gdc7-dmcn.csv?\
submission_year=2024&\
analysis_neighborhood='Lakeshore'"
Table.read_table(Lakeshore_2024_url)
```

---

### Throttling 

Many modern data sets are very large and update pretty frequently. Every request made with an API is associated with a cost. Some platforms limit your access to a few requests per day if you don't pay for more access. Another way platforms reduce costs is to limit the amount of data that is transferred by default with an API request. In other words, they throttle the data. DataSF is configured to only provide 1,000 rows of data in a single request unless you specify otherwise. The `$limit=80000` API parameter can be added to the URL to get **all** the available 2024 data. 

```
"https://data.sfgov.org/resource/gdc7-dmcn.csv?submission_year=2024&analysis_neighborhood='Lakeshore'&$limit=80000"
```

Having all of that data in your possession can be really helpful, but in many cases that file is really large, and the data changes frequently, so you end up running out of memory on your computer or end up working with old data.

---

### Task 09 💻

Using the DataSF API, get the 2024 Mission rental data where the base rent price includes electricity. Assign the data as a table to the name `Mission_2024`.

**Note**: There should only be a few hundred rentals that match this description, so you don't need to worry about throttling. 

```python
Mission_2024_url = "https://data.sfgov.org/resource/gdc7-dmcn.csv?\
submission_year=2024&\
analysis_neighborhood='Mission'&\
base_rent_includes_electricity='Y'"
Mission_2024 = Table.read_table(Mission_2024_url)
Mission_2024
```

```python
grader.check("task_09")
```

Great work accessing the DataSF data through their API. This is a huge step. A lot of your favorite platforms like Spotify, IMDB, etc. have a API so that you can access the data in similar way. Go explore!

---

### Task 10 💻

How many one-bedroom rentals are there (reported) in the 2024 Mission Housing Inventory? Assign that value to `one_bedrooms_count`.

```python
one_bedrooms_count = Mission_2024.where('bedroom_count', 'One-Bedroom').num_rows
one_bedrooms_count 
```

```python
grader.check("task_10")
```

---

## Reflection

This activity introduced you to the fundamentals of accessing and working with stored data. Through hands-on exercises, you learned how to read different data formats, manipulate tables, and explore datasets. While there are countless ways to store and access data, becoming familiar with a few common file types and retrieval methods will make it easier to learn new ones using documentation and tutorials. In most projects, you'll encounter a limited set of data formats and preferred access methods, so focusing on these will help you build confidence and efficiency in working with real-world datasets.

---

## License

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a>.

<img src="./by-nc-sa.png" width=100px>

Tests for task_01 (points: None)
```
>>> isinstance(hail, Table)
True
```
failure_message: "❌ hail should be assigned to a Table."
success_message: "✅ hail is assigned to a Table."

```
>>> hail.column('latitude').sum() + hail.column('longitude').sum()
-107668.5
```
failure_message: "❌ It doesn't seem like you've assigned hail to the correct data."
success_message: "✅ It seems like you've assigned hail to the correct data."



Tests for task_02 (points: None)
```
>>> import numbers
>>> isinstance(hail_num_rows, numbers.Integral)
True
```
failure_message: "❌ hail_num_rows should be assigned to an integer."
success_message: "✅ hail_num_rows is assigned to an integer."

```
>>> hail_num_rows
1745
```
failure_message: "❌ It doesn't seem like you've assigned hail_num_rows to the correct number."
success_message: "✅ It seems like you've assigned hail_num_rows to the correct number."



Tests for task_03 (points: None)
```
>>> import numbers
>>> isinstance(average_hail_size, numbers.Real)
True
```
failure_message: "❌ average_hail_size should be assigned to a numerical value."
success_message: "✅ average_hail_size is assigned to a numerical value."

```
>>> average_hail_size
147.08882521489971
```
failure_message: "❌ It doesn't seem like you've assigned average_hail_size to the correct number."
success_message: "✅ It seems like you've assigned average_hail_size to the correct number."



Tests for task_04 (points: None)
```
>>> isinstance(KS_hail, Table)
True
```
failure_message: "❌ KS_hail should be assigned to a Table."
success_message: "✅ KS_hail is assigned to a Table."

```
>>> KS_hail.labels
('timestamp', 'size', 'latitude', 'longitude', 'location')
```
failure_message: "❌ KS_hail doesn't have the correct column labels."
success_message: "✅ KS_hail has the correct column labels"

```
>>> KS_hail
timestamp                      | size | latitude | longitude | location
2021-04-07 04:45:00.000000 UTC | 100  | 37.29    | -98.58    | MEDICINE LODGE
2021-04-07 01:34:00.000000 UTC | 125  | 38.29    | -98.99    | 2 NNW PAWNEE ROCK
2021-04-07 02:45:00.000000 UTC | 150  | 38.29    | -98.76    | 6 SSE GREAT BEND
2021-04-07 03:03:00.000000 UTC | 100  | 38.36    | -98.79    | 1 E GREAT BEND
2021-04-07 03:39:00.000000 UTC | 100  | 38.36    | -98.58    | ELLINWOOD
2021-05-16 23:02:00.000000 UTC | 125  | 38.52    | -98.78    | HOISINGTON
2021-09-15 01:10:00.000000 UTC | 200  | 38.36    | -98.76    | 1 SE GREAT BEND
2021-09-15 01:40:00.000000 UTC | 100  | 38.52    | -98.77    | HOISINGTON
2021-06-12 02:00:00.000000 UTC | 175  | 37.79    | -97.11    | BENTON
2021-07-10 20:17:00.000000 UTC | 175  | 37.61    | -96.61    | 5 NNE LATHAM
... (291 rows omitted)
```
failure_message: "❌ It doesn't seem like you've assigned KS_hail to the correct table."
success_message: "✅ It seems like you've assigned KS_hail to the correct table"



Tests for task_05 (points: None)
```
>>> isinstance(KS_hail_sorted, Table)
True
```
failure_message: "❌ KS_hail_sorted should be assigned to a table."
success_message: "✅ KS_hail_sorted is assigned to a table."

```
>>> KS_hail_sorted.column('size').sum()
42650
```
failure_message: "❌ KS_hail_sorted doesn't seem to have the correct data."
success_message: "✅ KS_hail_sorted seems to have the correct data."

```
>>> np.all(KS_hail_sorted.column('size') == np.sort(KS_hail_sorted.column('size'))[::-1])
True
```
failure_message: "❌ KS_hail_sorted doesn't seem to be sorted correctly."
success_message: "✅ KS_hail_sorted seems to be sorted correctly."



Tests for task_06 (points: None)
```
>>> import numbers
>>> isinstance(TX_lat, numbers.Real)
True
```
failure_message: "❌ TX_lat is not a numerical value."
success_message: "✅ TX_lat is a numerical value."

```
>>> import numbers
>>> isinstance(TX_long, numbers.Real)
True
```
failure_message: "❌ TX_long is not a numerical value."
success_message: "✅ TX_long is a numerical value."

```
>>> TX_lat
30.76
```
failure_message: "❌ TX_lat is not correct."
success_message: "✅ TX_lat is correct."

```
>>> TX_long
-96.75
```
failure_message: "❌ TX_long is not correct."
success_message: "✅ TX_long is correct."



Tests for task_08 (points: None)
```
>>> isinstance(rent_upload, Table)
True
```
failure_message: "❌ You should have assigned rent_upload to a table."
success_message: "✅ You've assigned rent_upload to a table."



Tests for task_09 (points: None)
```
>>> isinstance(Mission_2024, Table)
True
```
failure_message: "❌ Mission_2024 should be assigned to a table."
success_message: "✅ Mission_2024 is assigned to a table."

```
>>> Mission_2024.num_rows
347
```
failure_message: "❌ Mission_2024 does not have the correct number of rows."
success_message: "✅ Mission_2024 has the correct number of rows."

```
>>> Mission_2024.num_columns
28
```
failure_message: "❌ Mission_2024 does not have the correct number of columns."
success_message: "✅ Mission_2024 has the correct number of columns."

```
>>> Mission_2024.column('unique_id').sum()
8918569704665515528
```
failure_message: "❌ Mission_2024 doesn't seem to contain the correct data."
success_message: "✅ Mission_2024 seems to contain the correct data."



Tests for task_10 (points: None)
```
>>> import numbers
>>> isinstance(one_bedrooms_count, numbers.Real)
True
```
failure_message: "❌ one_bedrooms_count should be assigned to a number."
success_message: "✅ one_bedrooms_count is assigned to a number."

```
>>> one_bedrooms_count
52
```
failure_message: "❌ one_bedrooms_count is not correct."
success_message: "✅ one_bedrooms_count is correct."

